Human-specific lncRNAs contributed critically to human evolution by distinctly regulating gene expression

Author Response

The following is the authors’ response to the original reviews.

We thank the two reviewers very much for their careful review and valuable comments. Upon these comments, the following revisions have been made. First, we have performed a new analysis on human accelerated regions (HARs) recently reported by the Zoonomia Project. Second, we have presented more data on experimentally detected and computationally predicted DBSs of MALAT1, NEAT1, and MEG3. Third, we have added details on the RNA-seq data processing and subsequent differential expression testing to the Materials and Methods section. Fourth, we have clarified some details on the human ancestor sequence and the use of parameters and thresholds. Six new citations are added. In addition, we have also carefully polished the main text. We hope these revisions, together with the Responses-to-Reviewers, would help the reader better get the information from the paper.

eLife assessment

In this valuable manuscript, the authors attempt to examine the role of long non-coding RNAs (lncRNAs) in human evolution, through a set of population genetics and functional genomics analyses that leverage existing datasets and tools. Although the methods are at times inadequate - for example, suitable methods and/or relevant controls are lacking at many points, and selection is inferred sometimes too quickly - the results nonetheless point towards a possible contribution of long non-coding RNAs to the evolution of human biology and they suggest clear directions for future, more rigorous study.

Public Reviews:

Reviewer #1 (Public Review):

Summary

While DNA sequence divergence, differential expression, and differential methylation analysis have been conducted between humans and the great apes to study changes that "make us human", the role of lncRNAs and their impact on the human genome and biology has not been fully explored. In this study, the authors computationally predict HSlncRNAs as well as their DNA Binding sites using a method they have developed previously and then examine these predicted regions with different types of enrichment analyses. Broadly, the analysis is straightforward and after identifying these regions/HSlncRNAs the authors examined their effects using different external datasets.

Strengths/weaknesses

By and large, the analysis performed is dependent on their ability to identify HSlncRNAs and their DBS. I think that they have done a good job of showing the performance metrics of their methods in previous publications. Thereafter, they perform a series of enrichment-type analyses that have been used in the field for quite a while now to look at tissue-specific enrichment, or region-specific enrichment, or functional enrichment, and I think these have been carried out well. The authors achieved the aims of their work. I think one of the biggest contributions that this paper brings to the field is their annotation of these HSlncRNAs. Thus a major revisionary effort could be spent on applying their method to the latest genomes that have been released so that the community could get a clean annotation of newly identified HSlncRNAs (see comment 2).

Comments

1. Though some of their results about certain HSlncRNAs having DBSs in all genes is rather surprising/suspicious, I think that broadly their process to identify and validate DBSs is robust, they have multiple lines of checks to identify such regions, including functional validation. These predictions are bound to have some level of false positive/negative rate and it might be nice to restate those here and on what experiment/validation data these were conducted. However, the rest of their analysis comprises different types of enrichment analysis which shouldn't be affected by outlier HSlncRNAs if indeed their FPR/FNR are low.

1. There are now several new genomes available as part of the Zoonomia consortium and 240 Primate consortium papers released. These papers have re-examined some annotations such as Human Accelerated Regions (HARs) and found with a larger dataset as well as better reference genomes, that a large fraction of HARs were actually incorrectly annotated - that is that they were also seen in other lineages outside of just the great apes. If these papers have not already examined HSlncRNAs, the authors should try and re-run the computational predictions with this updated set and then identify HSlncRNAs there. This might help to clarify their signal and remove lncRNAs that might be present in other primates but are somehow missing in the great apes. This might also help to mitigate some results that they see in section 3 of their paper in comparing DBS distances between archaics and humans.

Responses:

(1) Thanks for the good suggestion. We have checked the Zoonomia reported genomes and found that new primate genomes are monkeys and lemurs but not apes (Zoonomia Consortium. Nature 2023. https://doi.org/10.1038/s41586-020-2876-6), and the phylogenetic relationships between monkeys and humans are much more remote than those between apes and humans. In addition, the Zoonomia project did target identifying new lncRNA genes.

(2) We have examined the Zoonomia-reported HARs (Keough et al. Science 2023. DOI: 10.1126/science.abm1696). Of the 312 HARs reported by Keough et al, 8 overlap 26 DBSs of 14 HS lncRNAs; moreover, DBSs greatly outnumber HARs, suggesting that HAR and DBS are different sequences with different functions.

(3) In the revised manuscript, a new paragraph (the second one) has been added to the section “HS lncRNAs regulate diverse genes and transcripts” to describe the HAR analysis result.

1. The differences between the archaic hominins in their DBS distances to modern humans are a bit concerning. At some level, we expect these to be roughly similar when examining African modern humans and perhaps the Denisovan being larger when examining Europeans and Asians, but they seem to have distances that aren't expected given the demography. In addition, from their text for section 3, they begin by stating that they are computing two types of distances but then I lost track of which distance they were discussing in paragraph 3 of section 3. Explicitly stating which of the two distances in the text would be helpful for the reader.

Responses:

(1) Upon the archaic human genomes, the genomic distances from the three modern humans are shorter to Denisovan than to Altai Neanderthal; however, upon the related studies we cite, the phylogenetic relationship between the three modern humans is more remote to Denisovan than to Altai Neanderthal. Thus, the finding that 2514 and 1256 DBSs have distances >0.034 in Denisovans and Altai Neanderthals is not unreasonable. The numbers of DBSs, of course, depend on the cutoff of 0.034, which is somewhat subjective but not unreasonable.

(2) The second paragraph is added to the Discussion, discussing parameters and cutoffs.

(3) Regarding the two types of distance, the distances computed in the first way were not further analyzed because, as we note, “This anomaly may be caused by that the human ancestor was built using six primates without archaic humans”.

1. Isn't the correct control to examine whether eQTLs are more enriched in HSlncRNA DBSs a set of transcription factor binding sites? I don't think using just promoter regions is a reasonable control here. This does not take away from the broader point however that eQTLs are found in DBSs and I think they can perform this alternate test.

Responses:

Indeed, TFBSs are more comparable to DBSs than promoters. However, many more methods have been developed to predict TFBSs than to predict DBSs, making us concerned about TFBS prediction's reliability. Since most QTLs in DBSs are mQTLs (Supplementary Table 13), but many QTLs in TFBSs are eQTLs (Flynn et al. PLoS Genetics 2021. DOI: 10.1371/journal.pgen.1009719), it is pretty safe to conclude that DBSs are enriched in mQTLs.

1. In the Discussion, they highlight the evolution of sugar intake, which I'm not sure is appropriate. This comes not from GO enrichment but rather from a few genes that are found at the tail of their distribution. While these signals may be real, the evolution of traits is often highly polygenic and they don't see this signal in their functional enrichment. I suggest removing that line. Moreover, HSlncRNAs are ones that are unique across a much longer time frame than the transition to agriculture which is when sugar intake rose greatly. Thus, it's unlikely to see enrichment for something that arose in the past 6000-7000 years would in the annotation that is designed to detect human-chimp or human-neanderthal level divergence.

Responses:

(1) The Discussion on human adaptation to high sugar intake is based on both enriched GO terms (Supplementary Table 4, 7) and a set of genes in modern humans with the most SNP-rich DBSs (Table 2). These glucose-related GO terms are not at the tail of the list because, of the 614 enriched GO terms (enriched in genes with strongest DBSs), glucose metabolism-related ones are ranked 208, 212, 246, 264, 504, 522, 591, and of the 409 enriched GO terms (enriched in the top 1256 genes in Altai Neanderthals), glucose metabolism-related ones are ranked 152 and 217.

(2) Indeed, there are other top-ranked enriched GO terms; some (e.g., neuron projection development (GO:0031175) and cell projection morphogenesis (GO:0048858)) have known impact on human evolution, but the impact of others (e.g., cell junction organization (GO:0034330)) remain unclear. We specifically report human adaptation to high sugar intake because the DBSs in related genes show differences in modern humans (Table 2).

Reviewer #2 (Public Review):

Lin et al attempt to examine the role of lncRNAs in human evolution in this manuscript. They apply a suite of population genetics and functional genomics analyses that leverage existing data sets and public tools, some of which were previously built by the authors, who clearly have experience with lncRNA binding prediction. However, I worry that there is a lack of suitable methods and/or relevant controls at many points and that the interpretation is too quick to infer selection. While I don't doubt that lnc RNAs contribute to the evolution of modern humans, and certainly agree that this is a question worth asking, I think this paper would benefit from a more rigorous approach to tackling it.

At this point, my suggestions are mostly focused on tightening and strengthening the methods; it is hard for me to predict the consequence of these changes on the results or their interpretation, but as a general rule I also encourage the authors to not over-interpret their conclusions in terms of what phenotype was selected for when as they do at certain points (eg glucose metabolism).

Responses:

(1) Now, we use more cautious wording to describe the results.

(2) A paragraph (the second one) is added to Discussion to explain parameters and cutoffs.

(3) We make the caution at the end of the third paragraph that “We note that these are findings instead of conclusions, and they indicate, suggest, or support something revealing the primary question of what genomic differences critically determine the phenotypic differences between humans and apes and between modern and archaic humans”.

I note some specific points that I think would benefit from more rigorous approaches, and suggest possible ways forward for these.

1. Much of this work is focused on comparing DNA binding domains in human-unique long-noncoding RNAs and DNA binding sites across the promoters of genes in the human genome, and I think the authors can afford to be a bit more methodical/selective in their processing and filtering steps here. The article begins by searching for orthologues of human lncRNAs to arrive at a set of 66 human-specific lncRNAs, which are then characterised further through the rest of the manuscript. Line 99 describes a binding affinity metric used to separate strong DBS from weak DBS; the methods (line 432) describe this as being the product of the DBS or lncRNA length times the average Identity of the underlying TTSs. This multiplication, in fact, undoes the standardising value of averaging and introduces a clear relationship between the length of a region being tested and its overall score, which in turn is likely to bias all downstream inference, since a long lncRNA with poor average affinity can end up with a higher score than a short one with higher average affinity, and it's not quite clear to me what the biological interpretation of that should be. Why was this metric defined in this way?

Responses:

(1) Binding affinity and length of all DBSs of HS lncRNAs are given in Supplementary Table 2 and 3. Since a triplex (say, 100 bp in length) may have 50% or 70% of nucleotides bound, it is necessary to differentiate binding affinity and length, and the two measures can differentiate DBSs of the same length but with different binding affinity and DBSs with the same binding affinity but different length.

(2) Differentiating DBSs into strong and weak ones is somewhat subjective, accurately differentiating them demands experimental data that are currently unavailable, and it is advisable to separately analyze strong and weak DBSs because they may likely influence different aspects of human evolution.

1. There is also a strong assumption that identified sites will always be bound (line 100), which I disagree is well-supported by additional evidence (lines 109-125). The authors show that predicted NEAT1 and MALAT1 DBS overlap experimentally validated sites for NEAT1, MALAT1, and MEG3, but this is not done systematically, or genome-wide, so it's hard to know if the examples shown are representative, or a best-case scenario.

Responses:

(1) We do not assume/think that identified sites will always be bound. Instead, lncRNA/DBS binding is highly context-dependent (including tissue-specific).

(2) An extra supplementary table (Supplementary Table 15) is added to show what predicted DBSs overlap experimentally detected DBSs for NEAT1, MALAT1, and MEG3. By the way, it is more accurate to say “experimentally detected” than “experimentally validated”, because experimental data have true/false positives and true/false negatives, and different sequencing protocols (for detecting lncRNA/DNA binding) may generate somewhat different results.

It's also not quite clear how overlapping promoters or TSS are treated - are these collapsed into a single instance when calculating genome-wide significance? If, eg, a gene has five isoforms, and these differ in the 3' UTR but their promoter region contains a DBS, is this counted five times, or one? Since the interaction between the lncRNA and the DBS happens at the DNA level, it seems like not correcting for this uneven distribution of transcripts is likely to skew results, especially when testing against genome-wide distributions, eg in the results presented in sections 5 and 6. I do not think that comparing genes and transcripts putatively bound by the 40 HS lncRNAs to a random draw of 10,000 lncRNA/gene pairs drawn from the remaining ~13500 lncRNAs that are not HS is a fair comparison. Rather, it would be better to do many draws of 40 non-HS lncRNAs and determine an empirical null distribution that way, if possible actively controlling for the overall number of transcripts (also see the following point).

Responses:

(1) We analyzed each and every GENCODE-annotated transcript (Supplementary Table 2). For example, if a gene has N TSS and N transcripts, DBSs are predicted in N promoter regions. When analyzing gene expression in tissues, each and every transcript is analyzed.

(2) Ideally, it would be better to do many draws, but statistically, a huge number is needed due to the number of total genes in the human genome.

(3) We feel that doing many draws of 40 non-HS lncRNAs and determining an empirical null distribution is not as straightforward as comparing HS lncRNA-target transcript pairs (45% show significant expression correlation) with random lncRNA-random transcript pairs (2.3% show significant expression correlation).

1. Thresholds for statistical testing are not consistent, or always well justified. For instance, in line 142 GO testing is performed on the top 2000 genes (according to different rankings), but there's no description of the background regions used as controls anywhere, or of why 2000 genes were chosen as a good number to test? Why not 1000, or 500? Are the results overall robust to these (and other) thresholds? Then line 190 the threshold for downstream testing is now the top 20% of genes, etc. I am not opposed to different thresholds in principle, but they should be justified.

Responses:

(1) The over-representation analysis using g:Profiler was applied to the top and bottom 2000 genes with the whole genome as the background. The number “2000” was chosen somewhat subjectively. If more or fewer genes were chosen, more or fewer enriched GO terms would be identified, but GO terms with adjusted P-values <0.05 would be quite stable.

(2) A paragraph (the second one) is added to the Discussion to explain parameters and cutoffs.

Likewise, comparing Tajima's D values near promoters to genome-wide values is unfair, because promoters are known to be under strong evolutionary constraints relative to background regions; as such it is not surprising that the results of this comparison are significant. A fairer comparison would attempt to better match controls (eg to promoters without HS lncRNA DBS, which I realise may be nearly impossible), or generate empirical p-values via permutation or simulation.

Responses:

We examined Tajima’s D in DBSs (Supplementary Figure 9) and in HS lncRNA genes (Supplementary Figure 18). We compared the Tajima’s D values with the genome-wide background in both cases.

1. There are huge differences in the comparisons between the Vindija and Altai Neanderthal genomes that to me suggest some sort of technical bias or the such is at play here. e.g. line 190 reports 1256 genes to have a high distance between the Altai Neanderthal and modern humans, but only 134 Vindija genes reach the same cutoff of 0.034. The temporal separation between the two specimens does not seem sufficient to explain this difference, nor the difference between the Altai Denisovan and Neanderthal results (2514 genes for Denisovan), which makes me wonder if it is a technical artefact relating to the quality of the genome builds? It would be worth checking.

Responses:

(1) The cutoff of 0.034 was chosen upon that DBSs in the top 20% (4248) genes in chimpanzees have distances larger than this cutoff, and accordingly, 4248, 1256, 2514, and 134 genes have DBSs distances >0.034 in chimpanzees, Altai Neanderthals, Denisovans, and Vindija Neanderthals. These numbers of genes qualitatively agree with the phylogenetic distances from chimpanzees, archaic humans to modern humans. If a percentage larger or smaller than 20% (e.g., 10% or 30%) is chosen, and so is a cutoff X, the numbers of genes with DBSs distance >X would not be 4248, 1256, 2514, and 134, but could still qualitatively agree with the phylogenetic distances from chimpanzees, archaic humans to modern humans.

(2) The second paragraph in the Discussion now explains the parameters and cutoffs.

1. Inferring evolution: There are some points of the manuscript where the authors are quick to infer positive selection. I would caution that GTEx contains a lot of different brain tissues, thus finding a brain eQTL is a lot easier than finding a liver eQTL, just because there are more opportunities for it. Likewise, claims in the text and in Tables 1 and 2 about the evolutionary pressures underlying specific genes should be more carefully stated. The same is true when the authors observe high Fst between groups (line 515), which is only one possible cause of high Fst - population differentiation and drift are just as capable of giving rise to it, especially at small sample sizes.

Responses:

(1) We analyzed brain tissues separately instead of taking the whole brain as a tissue, see Supplementary Table 12 and Figure 3.

(2) We make the caution at the end of the third paragraph that “We note that these are findings instead of conclusions, and they indicate, suggest, or support something revealing the primary question of what genomic differences critically determine the phenotypic differences between humans and apes and between modern and archaic humans”.

Reviewer #1 (Recommendations For The Authors):

Some figures are impossible to see/read so I wasn't able to evaluate them - Fig, 1B, 1E, 1F are small and blurry.

Responses:

High-quality figures are provided.

Typo in line 178: in these archaic humans, the distances of HS lncRNAs are smaller than the distances of DBSs.

Responses:

This is not a typo. We use “distance per base” to measure whether HS lncRNAs or their DBSs have evolved more from archaic humans to modern humans. See also Supplementary Note 4 and 5.

Reviewer #2 (Recommendations For The Authors):

1. There's some inconsistency in the genome builds and the database versions used, eg, sometimes panTro4 is used and sometimes panTro5 (line 456). Likewise, the version of GENCODE used is very old (18), the current version is 43. The current version contains 19928 lncRNAs, which is a big difference relative to what is being tested!

Responses:

(1) panTro4 was used to search orthologues of human lncRNAs; this time-consuming work started several years ago when the version of GENCODE was V18 (see Lin et al., 2019).

(2) Regarding “the version of GENCODE used is very old (V18)”, we have later examined the 4396 human lncRNAs reported in GENCODE V36 and found that the set of 66 HS lncRNAs remains the same.

(3) The counterparts of HS lncRNAs’ DBSs in chimpanzees were predicted recently using panTro5.

1. Table 1: What does 'mostly' mean in this context? I understand that it refers to sequence differences between humans and the other genomes, but what is the actual threshold, and how is it defined?

Responses:

The title of Table 1 is “Genes with strongest DBSs and mostly changed sequence distances from modern humans to archaic humans and chimpanzees”. Instead of using two cutoffs, choosing genes with the two features seems easy and sensible.

1. Line 117: The methods do not include information on the RNA-seq data processing and subsequent DE testing.

Responses:

The details are added to the section “Experimentally validating DBS predictiom” (The reads were aligned to the human GRCh38 genome using Hiasat2 (Kim et al., 2019), and the resulting sam files were converted to bam files using Samtools (Li et al., 2009). Stringtie was used to quantify gene expression level (Pertea et al., 2015). Fold change of gene expression was computed using the edgeR package (Robinson et al., 2010), and significant up- and down-regulation of target genes after DBD knockout was determined upon |log2(fold change)| > 1 with FDR < 0.1).

1. Line 180: I looked at the EPO alignment and it's not clear to me what 'human ancestor' means, but it may well explain the issues the authors have with calculating distances (I agree those numbers are weird). Is it the reconstructed ancestral state of humans at around 300-200,000 years ago (coalescence of most human uniparental lineages), or the inferred sequence of the human-chimpanzee most recent common ancestor? If it's the former, it's not surprising it skews results towards shorter distances for modern humans, since the tree distance from that point to archaic hominins is significantly larger than to modern humans.

Responses:

The “human ancestor” is constructed by the EBI team upon the genomes of six primates in the Ensembl website. We find that the reconstructed ancestral state of humans may be unlikely around 300,000-200,000 years, and may be much earlier. We also find that many DNA sequences of the “human ancestor” are low-confidence calls (i.e., the ancestral states are supported by only one primate’s sequence).

1. Line 221: SNP-rich DBS: Is this claim controlled for the length of the DBS?

Responses:

No. Long DBSs tend to have more SNPs. When comparing the same DBS in modern humans, archaic humans, and chimpanzees, both the length and SNP number reflect evolution, so it is not necessary to control for the length.

1. Given that GTEx is primarily built off short-read data and it is impossible to link binding of a lncRNA to a DBS with its impact with a specific transcript

Responses:

As written in the section “Examining the tissue-specific impact of HS lncRNA-regulated gene expression”, we calculated the pairwise Spearman's correlation coefficient between the expression of an HS lncRNA (the representative transcript, median TPM value > 0.1) and the expression of each of its target transcripts (median TPM value > 0.1) using the scipy.stats.spearmanr program in the scipy package. The expression of an HS lncRNA gene and a target transcript was considered to be significantly correlated if the |Spearman's rho| > 0.3, with Benjamini-Hochberg FDR < 0.05.

1. Line 429: should TTO be TFO?

Responses:

Here TTO should be TFO; the typo is corrected.

1. Methods, section 7: Some of the text in this section should perhaps be moved to the results section?

Responses:

Each of the two paragraphs in Methods’ section 7 is quite large, and some contents in Supplementary Notes are also very relevant. Thus, moving them to the Results section could make the Results too lengthy and specific.

1. Line 587: GTEx is built from samples of primarily European ancestry and has poor representation of African ancestry and negligible representation of Asian ancestry (see the GTEx v8 paper supplement). This means that it is basically impossible to find a non-European population-specific eQTL in GTEx, which in turn impacts these results.

Responses:

(1) Indeed, this is a serious issue of data analysis, and this issue cannot be solved until more Africans are sequenced.

(2) Anyway, one can still find considerable African-specific eQTLs in GTEx, such as rs28540058 (with frequency of 0, 0, 0.13 in CEU, CHB, YRI) and rs58772997 (with frequency of 0, 0, 0.12 in CEU, CHB, YRI (see Supplementary Table12 and Supplementary Figure 22).

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: What genomic sequences make protein-coding genes generate divergent expression in closely related species, specifically, differentiate humans from apes, puzzle many researchers. Many studies examined species-specific gene birth, gene loss, and changes in promoters and transcription factor binding sites, but the identification and impact of human-specific lncRNAs remain unexplored. This study identified human-specific lncRNAs from GENCODE-annotated human lncRNAs, predicted their DNA binding sites (DBSs) genome-wide, and analyzed the DBSs and their counterparts in modern humans (CEU, CHB, and YRI), archaic humans (Altai Neanderthals, Denisovans, and Vindija Neanderthals), and chimpanzees. The results reveal how human-specific lncRNAs and their DBSs have transcriptionally regulated gene expression human-specifically. The rewiring of gene expression has undergone continuous evolution, significantly changed gene expression in the brain, promoted the adaptive evolution of humans, and influenced differences in modern humans. These results reveal the importance of human-specific lncRNAs (for human evolution) and highlight the importance of other species-specific lncRNAs.

Revised version: What genomic sequences make conserved genes generate divergent expression in closely related species, which may have critically driven human evolution, has puzzled researchers for decades. Genomic studies have examined species-specific gene birth, gene loss, and changes in promoters and transcription factor binding sites, but species-specific epigenetic regulation remains barely explored. This study identified human-specific long noncoding RNAs (lncRNAs) from GENCODE-annotated human lncRNAs, predicted their DNA binding sites (DBSs) genome-wide, analyzed these DBSs and their counterparts in modern humans (CEU, CHB, and YRI), archaic humans (Altai Neanderthals, Denisovans, and Vindija Neanderthals), and chimpanzees, and analyzed the impact of DBSs on gene expression in modern and archaic humans. The results suggest that human-specific lncRNAs and their DBSs have substantially rewired gene expression human-specifically and that the rewiring has evolved continuously from archaic to modern humans. Rewired gene expression promotes brain development, makes humans adapt to new environments and lifestyles, and causes differences in modern humans. These results uncover a critical dimension of human evolution and underscore the diverse functions of species-specific lncRNAs.

Introduction

Earlier version: Genomic changes that drive evolution include gene birth (Kaessmann, 2010), gene loss (Albalat and Canestro, 2016), and changes in regulatory sequences (Prud’homme et al., 2007). Small genomic differences between humans and other hominids, but significant phenotypic and behavioral differences, make what genomic changes at what time critically drove human evolution an enduring puzzle (Anton et al., 2014; Pollen et al., 2023). Waves of studies have identified genes that promote human brain development (Evans et al., 2005; Fiddes et al., 2018; Mekel-Bobrov et al., 2005; Pinson et al., 2022; Suzuki et al., 2018). However, conclusions from gene birth-centered studies are controversial (Currat et al., 2006; Timpson et al., 2007; Yu et al., 2007). In particular, genes important for brain development (e.g., NOTCH2NL and ASPM) do not necessarily account for many human traits that preceded or coincided with brain enlargement. Recent studies revealed the importance of regulatory sequences for human evolution, including new sequences (Liu et al., 2021), lost sequences (McLean et al., 2011), sequences showing accelerated evolution (Dong et al., 2016; Mangan et al., 2022; Prabhakar et al., 2006; Whalen and Pollard, 2022), and the turnover of transcription factor binding sites (TFBS) (Krieger et al., 2022; Otto et al., 2009; Zhang et al., 2023). An unexplored question is how simian- and human-specific lncRNAs, with the sequences of their DNA binding sites, regulate conserved genes and drive human evolution.

Revised version: Genomic changes that drive evolution include gene birth (Kaessmann, 2010), gene loss (Albalat and Canestro, 2016), and changes in regulatory sequences (Prud’homme et al., 2007). The limited genomic differences but substantial phenotypic and behavioral divergence between humans and other hominids make what genomic changes at what time have critically driven human evolution an enduring puzzle (Anton et al., 2014; Pollen et al., 2023). Although waves of studies have identified genes promoting human brain development (Evans et al., 2005; Fiddes et al., 2018; Mekel-Bobrov et al., 2005; Pinson et al., 2022; Suzuki et al., 2018), conclusions from gene birth-centered studies are controversial (Currat et al., 2006; Timpson et al., 2007; Yu et al., 2007). Moreover, genes important for brain development (e.g., NOTCH2NL and ASPM) do not necessarily regulate other human traits that precede or coincide with brain enlargement. Recent studies revealed the importance of regulatory sequences for human evolution, including new sequences (Liu et al., 2021), lost sequences (McLean et al., 2011), human accelerated regions (HARs) (Dong et al., 2016; Mangan et al., 2022; Prabhakar et al., 2006; Whalen and Pollard, 2022), and the turnover of transcription factor binding sites (TFBS) (Krieger et al., 2022; Otto et al., 2009; Zhang et al., 2023). An unexplored question is how epigenetic regulation contributes to human evolution.

Earlier version: On this question, experimental studies have generated multiple findings. (a) The chimpanzee and human genomes are almost 99% identical (Chimpanzee Sequencing and Analysis Consortium, 2005). (b) About one-third of human lncRNAs are primate-specific (Derrien et al., 2012). (c) Species-specific lncRNAs have distinct expression profiles in organs of different species (Sarropoulos et al., 2019). (d) Many lncRNAs can bind to DNA sequences by forming RNA:DNA triplexes (Abu Almakarem et al., 2012), thereby recruiting histone and DNA modification enzymes to binding sites to regulate nearby gene expression (Lee, 2009; Roadmap Epigenomics et al., 2015). (e) Approximately 40% of differentially expressed genes in humans result from interspecies epigenetic differences (Hernando-Herraez et al., 2015). Compared with gene-centered hypotheses of human evolution, these findings suggest an alternative hypothesis – distinctly rewired gene expression by human-specific (HS) lncRNAs and their DNA binding sites has critically promoted human evolution and determined human traits.

Revised version: Experimental studies have generated substantial findings on epigenetic regulation. (a) The chimpanzee and human genomes are almost 99% identical (Chimpanzee Sequencing and Analysis Consortium, 2005). (b) About one-third of human long-noncoding RNAs (lncRNA) are primate-specific (Derrien et al., 2012). (c) Species-specific lncRNAs have distinct expression profiles in organs of different species (Sarropoulos et al., 2019). (d) Many lncRNAs can bind to DNA sequences by forming RNA:DNA triplexes (Abu Almakarem et al., 2012), thereby recruiting histone and DNA modification enzymes to binding sites to regulate nearby gene expression (Lee, 2009; Roadmap Epigenomics et al., 2015). (e) Approximately 40% of differentially expressed human genes result from interspecies epigenetic differences (Hernando-Herraez et al., 2015). Compared with protein-coding gene-centered hypotheses of human evolution, the alternative one is plausible – human-specific (HS) lncRNAs and their DNA binding sites have human-specifically rewired gene expression to critically promote human evolution and determine human traits.

Earlier version: Orthologous lncRNA genes can be identified by sequence search by structure and sequence alignment (Nawrocki, 2014; Nawrocki and Eddy, 2013). HS lncRNAs can be identified upon orthologues of human lncRNAs in other mammals. RNA:DNA triplexes follow specific base-pairing rules (Abu Almakarem et al., 2012). These rules allow the prediction of lncRNAs’ DNA binding domains (DBD) and DNA binding sites (DBS) (DBDs and DBSs indicate predicted sequences of DNA binding domains and binding sites) (Lin et al., 2019; Wen et al., 2022). Gene expression in most human organs and tissues has been profiled (GTEx Consortium, 2017). Genomes of modern humans, archaic humans, and multiple apes have been sequenced (1000 Genomes Project Consortium et al., 2012; Chimpanzee Sequencing and Analysis Consortium, 2005; Meyer et al., 2012; Prufer et al., 2017; Pru fer et al., 2013). These results enable us to examine the above alternative hypothesis.

Revised version: RNA sequence search by structure and sequence alignment can identify orthologous lncRNA genes in genomes (Nawrocki, 2014; Nawrocki and Eddy, 2013). Thus, HS lncRNAs can be identified upon human lncRNAs’ orthologues in other mammals. Many lncRNA can bind to DNA sequences following specific base-pairing rules to form RNA:DNA triplexes (Abu Almakarem et al., 2012). Thus, these rules allow computational prediction of lncRNAs’ DNA binding domains (DBD) and DNA binding sites (DBS) (Lin et al., 2019; Wen et al., 2022). Gene expression in most human organs and tissues has been profiled (GTEx Consortium, 2017). Further, cenomes of modern humans, archaic humans, and multiple apes have been sequenced (1000 Genomes Project Consortium et al., 2012; Chimpanzee Sequencing and Analysis Consortium, 2005; Meyer et al., 2012; Prufer et al., 2017; Pru fer et al., 2013). Together, these results make it feasible to explore the above alternative hypothesis.

Earlier version: HS lncRNAs and their DBSs may have occurred in three situations – (a) HS lncRNAs occurred before their DBSs, (b) HS lncRNAs and their DBSs occurred simultaneously, (c) HS lncRNAs occurred after their DBSs. It is also unknown which situation the rewiring hypothesis corresponds to. In this study, we identified HS lncRNAs upon the orthologues of the GENCODE-annotated human lncRNA genes in 16 mammalian genomes and predicted HS lncRNAs’ DBDs and DBSs in the genome of modern humans (CEU, CHB, and YRI) and these DBSs’ counterparts in the genomes of archaic humans (Altai Neanderthals, Denisovans, Vindija Neanderthals), and chimpanzees (Figure 1A). Then, we analyzed the sequences and evolution of HS lncRNAs and their DBSs, functional enrichment of target genes, and impact of HS lncRNA regulation on gene expression in GTEx organs and tissues in archaic and modern humans. Our genome and transcriptome analyses suggest that HS lncRNAs and their DBSs have substantially and distinctly rewired gene expression during human evolution. The rewiring of gene expression also greatly determines differences between human populations.

Revised version: In this study, we identified HS lncRNAs upon the orthologues of the GENCODE-annotated human lncRNA genes in 16 mammalian genomes, predicted HS lncRNAs’ DBDs and DBSs in the genomes of modern humans (CEU, CHB, and YRI), predicted these DBSs’ counterparts in the genomes of archaic humans (Altai Neanderthals, Denisovans, Vindija Neanderthals) and chimpanzees (Figure 1A), and experimentally validated DBS prediction. Then, we applied genome analyses and evolutionary/population genetics analyses to HS lncRNAs and their DBSs, applied transcriptome analysis to modern and archaic humans using the GTEx project data, and analyzed multiple experimental datasets. Our genome, transcriptome, and evolutionary analyses suggest that HS lncRNAs and their DBSs have continuously and distinctly rewired gene expression during human evolution. The rewired gene expression influences or determines diverse human traits, from an enlarged brain to modified metabolism pathways.

Earlier version: Study Overview. (A) The relationships between the chimpanzees, the three archaic humans (Altai Neanderthals, Denisovans, and Vindija Neanderthals), and modern humans. The dashed lines indicate the phylogenetic distances from modern humans to archaic humans and chimpanzees. The genes and DBSs have the following implications – the DBS in B2M lacks a counterpart in chimpanzees; the DBS in ABL2 has great differences in archaic humans; the DBS in IRNAR1 is highly polymorphic in modern humans. Red characters in DBSs indicate tissue-specific eQTLs or population-specific favored mutations. (B) The distributions of DBSs’ length and binding affinity. (C) The distributions of the numbers of HS lncRNAs’ target genes and transcripts. (D) The targeting relationships between HS lncRNAs. (E) The distributions of sequence distances of DBSs from modern humans to chimpanzees and archaic humans. (F) The impacts of HS lncRNA regulation on gene expression in GTEx tissues.

Revised version: Study Overview. (A) The relationships between the chimpanzees, the three archaic humans (Altai Neanderthals, Denisovans, and Vindija Neanderthals), and the three modern human populations. Dashed lines indicate the phylogenetic distances from modern humans to archaic humans and chimpanzees. Genes and DBSs have the following implications – the DBS in B2M lacks a counterpart in chimpanzees; the DBS in ABL2 has great differences between archaic and modern humans; the DBS in IRNAR1 is highly polymorphic in modern humans. Red letters in DBSs indicate tissue-specific eQTLs or population-specific favored mutations. (B) DBSs’ mean length and binding affinity. (C) The target genes and transcripts of 66 HS lncRNAs. (D) The targeting relationships between HS lncRNAs. (E) The sequence distances of DBSs (the top 40%) from modern humans to chimpanzees and archaic humans. (F) The impacts of interactions between HS lncRNAs and DBSs on gene expression in GTEx tissues.

Results

HS lncRNAs regulate diverse genes and transcripts

Earlier version: After searching for the orthologues of the 13562 GENCODE-annotated human lncRNA genes (v18) in 16 mammalian genomes by sequence and structure alignment using the Infernal program (Lin et al., 2019; Nawrocki, 2014; Nawrocki et al., 2009), we identified 66 HS lncRNA genes (Supplementary Table 1). We predicted HS lncRNAs’ DBSs in the 5000 bp promoter regions of the 179128 Ensembl-annotated transcripts (release 79) using the LongTarget program (Supplementary Note 1) (Lin et al., 2019). 105141 strong DBSs (binding affinity > 60, mean length > 147 bp) were identified in 96789 transcripts of 23396 genes, and 152836 weak DBSs (36 < binding affinity < 60) were identified in 127898 transcripts of 33185 genes (Figure 1B; Supplementary Table 2, 3). Genes and transcripts whose promoter regions contain the DBS of an HS lncRNA were assumed targets of the HS lncRNA. Several HS lncRNAs have DBSs in nearly all genes and transcripts (e.g., RP11-423H2.3 has strong DBSs in 73855 transcripts of 17941 genes), and only about 1.5% of target genes (0.6% of target transcripts) are human-specific (Figure 1C). Many transcripts have DBSs for multiple HS lncRNAs, many HS lncRNAs have multiple DBSs in a single transcript, and HS lncRNAs themselves show complex targeting relationships (Figure 1D; Supplementary Note 2). These findings suggest that HS lncRNAs substantially regulate gene expression.

Revised version: After searching for orthologues of the 13562 GENCODE-annotated human lncRNA genes (v18) in 16 mammalian genomes and the genomes of archaic humans by RNA sequence and structure alignment using the Infernal program (Figure 1A) (Lin et al., 2019; Nawrocki, 2014; Nawrocki et al., 2009), we identified 66 HS lncRNA genes (Supplementary Table 1). We predicted HS lncRNAs’ DBSs in the 5000 bp promoter regions of the 179128 Ensembl-annotated transcripts (release 79) using the LongTarget program (Supplementary Note 1) (Lin et al., 2019). 105141 strong DBSs (mean binding affinity >60, mean length >147 bp) were identified in 96789 transcripts of 23396 genes, and 152836 weak DBSs (36< mean binding affinity <60) were identified in 127898 transcripts of 33185 genes (Figure 1B; Supplementary Table 2, 3). Genes and transcripts whose promoter regions contain the DBS of an HS lncRNA were assumed targets of the HS lncRNA. Several HS lncRNAs have DBSs in nearly all genes and transcripts (e.g., RP11-423H2.3 has strong DBSs in 73855 transcripts of 17941 genes), and only about 1.5% of target genes (0.6% of target transcripts) are human-specific genes and transcripts (Figure 1C), indicating that HS lncRNAs substantially rewire conserved genes and supporting the above rewiring hypothesis. Many transcripts have DBSs for multiple HS lncRNAs, many HS lncRNAs have multiple DBSs in a single transcript, and HS lncRNAs themselves show complex targeting relationships (Figure 1D; Supplementary Note 2), indicating the complexity of the rewiring.

Added in the revised version: HARs are conserved genomic loci but have evolved acceleratedly in the human lineage. Multiple researchers assume that HARs critically influence human evolution (especially brain evolution); however, results obtained from HAR studies are controversial (Levchenko et al., 2018). The Zoonomia Project recently identified 312 HARs important for 3D genome rewiring and neurodevelopment (Keough et al., 2023). We found that only 8 HARs overlap 26 DBSs of 14 HS lncRNAs, suggesting that DBS and HAR are sequences of different functions subject to separate examination.

Earlier version: DBS prediction was validated and supported by multiple methods and datasets (Supplementary Note 1; Supplementary Fig 1-5). (a) DBSs of HS lncRNAs overlap well with experimentally identified DNA methylation and histone modification signals in multiple cell lines (https://genome.UCSC.edu). (b) DBSs of NEAT1 and MALAT1 overlap well with experimentally identified DNA binding sites of NEAT1, MALAT1, and MEG3 (Mondal et al., 2015; West et al., 2014). (c) Many DBSs co-localize with annotated cis-regulatory elements (cCREs) in genes’ promoter regions (https://genome.UCSC.edu). (d) We used CRISPR/Cas9 to knock out the DBD (100-200 bp) of seven lncRNAs (three HS lncRNAs and four wrongly transcribed lncRNAs) in multiple cell lines and performed RNA-seq before and after DBD knockout. Gene expression analysis revealed that the |fold change| of target genes was significantly larger than the |fold change| of non-target genes (one-sided Mann-Whitney test, p = 3.1e-72, 1.49e-114, and 1.12e-206 for the HS lncRNA RP13-516M14.1, RP11-426L16.8, and SNORA59B, and p = 2.58e-09, 6.49e-41, 0.034, and 5.23e-07 for the four wrongly transcribed lncRNAs). The results suggest that the changed expression of target genes was caused by the knockout of DBDs. (e) Finally, the LongTarget program uses the Smith-Waterman local alignment to simultaneously identify DBDs and DBSs. The alignment algorithm determines that a DBS of 147 bp is extremely unlikely to be generated by chance (p < 8.2e-19 to 1.5e-48).

Revised version: We validated DBS prediction using multiple methods and datasets (Supplementary Note 1; Supplementary Fig 1-5). (a) HS lncRNA DBSs overlap well with experimentally identified DNA methylation and histone modification signals in multiple cell lines (https://genome.UCSC.edu). (b) Many DBSs co-localize with annotated cis-regulatory elements (cCREs) in genes’ promoter regions (https://genome.UCSC.edu). (c) predicted and experimentally detected DBSs of NEAT1, MALAT1, MEG3 overlap well (Mondal et al., 2015; West et al., 2014). (d) We used CRISPR/Cas9 to knock out predicted DBDs (100-200 bp) of seven lncRNAs (three HS lncRNAs and four wrongly transcribed lncRNAs) in multiple cell lines, performed RNA-seq before and after DBD knockout, and analyzed differential gene expression The |fold change| of target genes (genes with DBS) was significantly larger than the |fold change| of non-target genes (genes without DBS) (one-sided Mann-Whitney test, p = 3.1e-72, 1.49e-114, and 1.12e-206 for the HS lncRNA RP13-516M14.1, RP11-426L16.8, and SNORA59B, and p = 2.58e-09, 6.49e-41, 0.034, and 5.23e-07 for the four wrongly transcribed lncRNAs), suggesting that changed expression of target genes is caused by DBD knockout. (e) Finally, the LongTarget program uses the Smith-Waterman local alignment to identify DBD/DBS binding; the alignment algorithm determines that a DBS of 147 bp is extremely unlikely to be generated by chance (p < 8.2e-19 to 1.5e-48).

Earlier version: Target genes characterize human evolution and traits

Revised version: HS lncRNAs’ target genes characterize human evolution and traits

Earlier version: Since strong and weak DBSs could occur early and late in human evolution, we examined whether their target genes were enriched in specific functions important for human evolution. We sorted target genes by DBS’s binding affinity, chose the top 2000 genes (DBS sequence > 252 bp, DBS binding affinity > 151) and bottom 2000 genes (DBS sequence < 60 bp, DBS binding affinity > 36), and performed over-representation analysis using the g:Profiler program and Gene Ontology (GO) database (Supplementary Table 4). The two target gene sets are co-enriched in GO terms important for human evolution, including “neuron projection development” (GO:0031175), “learning” (GO:0007612)), “regulation of anatomical structure size” (GO: GO:0090066), “head development” (GO:0060322), and “regulation of locomotion” (GO:0040012)) (Supplementary Note 3; Supplementary Table 4). 514 DBSs have binding affinity > 300 (DBS sequences > 500 bp) (Supplementary Table 5). Genes containing these strongest DBSs include IFNAR1 and NFATC1 (important for immune function), KIF21B and NTSR1 (critical for neural development), SLC2A11 and SLC2A1 (involved in glucose usage), BAIAP3 and TAS1R3 (a brain-specific angiogenesis inhibitor and a receptor for the sweet taste response), and some primate-specific noncoding RNA genes (e.g., CTD-3224I3.3 which is highly expressed in the cerebellum, lung, and testis) (Table 1). Interestingly, genes with weak DBSs are enriched in GO terms including “response to alcohol” (GO:0097305), “response to temperature stimulus” (GO:0009266), and “DNA methylation or demethylation” (GO:0044728), and “female pregnancy” (GO:0007565) (Supplementary Note 3; Supplementary Table 4).

Revised version: Since strong and weak DBSs may occur at different times in human evolution, we separately examined whether their target genes are enriched in functions important for human evolution. We sorted target genes by DBS binding affinity, chose the top 2000 genes (DBS length >252 bp and DBS binding affinity >151) and bottom 2000 genes (DBS length <60 bp but DBS binding affinity >36), and applied over-representation analysis using the g:Profiler program and Gene Ontology (GO) database to the two gene sets (Supplementary Table 4). The top and bottom 2000 genes are co-enriched in GO terms important for human evolution (instead of just brain evolution), including “neuron projection development” (GO:0031175), “learning” (GO:0007612), “regulation of anatomical structure size” (GO: GO:0090066), “head development” (GO:0060322), and “regulation of locomotion” (GO:0040012)) (Supplementary Note 3; Supplementary Table 4). 514 strongest DBSs have binding affinity >300 (DBS length >500 bp) (Supplementary Table 5), and genes containing these strongest DBSs include IFNAR1 and NFATC1 (important for immune function), KIF21B and NTSR1 (critical for neural development), SLC2A11 and SLC2A1 (involved in glucose usage), BAIAP3 and TAS1R3 (a brain-specific angiogenesis inhibitor and a receptor for the sweet taste response), and some primate-specific noncoding RNA genes (e.g., CTD-3224I3.3 which is highly expressed in the cerebellum, lung, and testis) (Table 1). Of note, genes with weak DBSs are enriched in GO terms including “response to alcohol” (GO:0097305), “response to temperature stimulus” (GO:0009266), “DNA methylation or demethylation” (GO:0044728), and “female pregnancy” (GO:0007565) (Supplementary Note 3; Supplementary Table 4). These findings help reveal how epigenetic regulation reshapes different human traits at different times.

The rewiring of gene expression has greatly evolved

Earlier version: The human genome is almost 99%, 99%, and 98% identical to the genomes of chimpanzees, bonobos, and gorillas (Chimpanzee Sequencing and Analysis Consortium, 2005; Prufer et al., 2012). Given these HS lncRNA genes, we assessed to what extent their DBS sequences are human-specific (whether DBSs occurred before, at the same time, or after HS lncRNAs). 97.81% of strong DBSs were found to have counterparts in chimpanzees, suggesting that most DBS sequences occurred before HS lncRNAs.

Revised version: The human genome is almost 99%, 99%, and 98% identical to the genomes of chimpanzees, bonobos, and gorillas (Chimpanzee Sequencing and Analysis Consortium, 2005; Prufer et al., 2012). Since where sequence differences are in these genomes is a question, we assessed to what extent HS lncRNAs’ DBS sequences are human-specific and found that 97.81% of strong DBSs have counterparts in chimpanzees. This suggests that many DBSs are not human-specific.

Earlier version: To examine the evolution of HS lncRNAs and their DBSs during human evolution, we identified their counterparts in Altai Neanderthals, Denisovans, and Vindija Neanderthals (Meyer et al., 2012; Prufer et al., 2017; Pru fer et al., 2013) (Figure 1A). HS lncRNAs and strong DBSs have counterparts in these archaic humans, but the distances of HS lncRNAs are small than the distances of DBSs between archaic and modern humans (all “distance” means distance per base). We computed DBS distances in two ways – from the human ancestor (downloaded from the EBI website) to chimpanzees, archaic humans, and modern humans (the first type of distances) and between the modern humans and chimpanzees and archaic humans (the second type of distances). In the first type of distances, many DBSs’ distances from the human ancestor to modern humans were smaller than to archaic humans. For example, DBSs in the top 20% (5033) genes in the human ancestor have distances > 0.015 to modern humans, but more DBSs have distances > 0.015 to Altai Neanderthals, Denisovans, and even chimpanzees (Supplementary Note 3). This defect may be caused by that the human ancestor was built using six primates without archaic humans. The second type of distances agree well with the phylogenetic distances between chimpanzees, archaic humans, and modern humans and suggest that greatly changed sequences occurred in very limited DBSs (Figure 1AE). For example, DBSs in the top 20% (4248) genes in chimpanzees have distances > 0.034, and DBSs having distances > 0.034 in genes in Altai Neanderthals, Denisovans, and Vindija Neanderthals are 1256, 2514, and 134 (Supplementary Table 6).

Revised version: To examine the evolution of HS lncRNAs and their DBSs during human evolution, we identified their counterparts in Altai Neanderthals, Denisovans, and Vindija Neanderthals (Meyer et al., 2012; Prufer et al., 2017; Prüfer et al., 2013). HS lncRNAs and strong DBSs have counterparts in these archaic humans, but the distances of HS lncRNAs are smaller than the distances of DBSs between archaic and modern humans (“distance” means distance per base), supporting that many DBSs occur in ape genomes instead of human-specific. We computed DBS distances in two ways. The first is from the reconstructed human ancestor (downloaded from the EBI website) to chimpanzees, archaic humans, and modern humans; the second is from chimpanzees and archaic humans to modern humans. The first type of distances has the anomaly that many DBSs’ distances from the human ancestor to modern humans are smaller than to archaic humans (for example, DBSs in the top 20% (5033) genes in the human ancestor have distances >0.015 to modern humans, but more DBSs have distances >0.015 to Altai Neanderthals, Denisovans, and even chimpanzees) (Supplementary Note 3). This anomaly may be caused by the human ancestor being built using six primates without archaic humans; therefore, no analysis was followed. In contrast, the second type of distances agrees well with the phylogenetic distances between chimpanzees, archaic humans, and modern humans. Notable, large sequence changes occurred in limited DBSs (Figure 1E), and genes containing these greatly changed DBSs may critically characterize human evolution. Given the high similarity between human and chimpanzee genomes, the top 20% (4248) genes in chimpanzees may sufficiently characterize the human-chimpanzee differences, and these genes have DBS distances >0.034. Taking this as a threshold, 1256, 2514, and 134 genes in Altai Neanderthals, Denisovans, and Vindija Neanderthals have DBS distance >0.034, and these numbers agree with the phylogenetic relationships between these archaic and modern humans (Figure 1A; Supplementary Table 6).

Earlier version: To examine whether the target genes of the mostly changed DBSs in chimpanzees and archaic humans are enriched in specific functions, we performed over-representation analysis for genes with using the g:Profiler program and Gene Ontology (GO) database. Since 4248, 1256, and 2514 genes in chimpanzees, Altai Neanderthals, and Denisovans have DBS distances > 0.034 (to modern humans), to make the statistical significance comparable, we examined the top 1256 genes in each analysis. The two sets of genes with largest DBS distances from modern humans to chimpanzees and Altai Neanderthals are enriched in many common GO terms, but “behavior” (GO:0007610) and “cognition” (GO:0050890) have higher significance in Altai Neanderthals than in chimpanzees (Figure 2). In addition, only genes in Altai Neanderthals are also enriched in such GO terms as “learning” (GO:0007612), “learning or memory” (GO:0007611), and “glucose homeostasis” (GO:0042593). The 134 genes whose DBS distances > 0.034 from Vindija Neanderthals to modern humans, as expected, are enriched only in two GO terms with low significance, including “smoothened signaling pathway” (GO:0007224) and “regulation of smoothened signaling pathway” (GO:0008589) (Supplementary Table 7). These results reveal that critical DBS changes occurred in genes involved in brain functions and some metabolism pathways in early archaic humans.

Revised version: To examine target genes with mostly changed DBSs are enriched in what functions, we applied over-representation analysis using the g:Profiler program and Gene Ontology (GO) database to target genes in chimpanzees, Altai Neanderthals, Denisovans, and Vindija Neanderthals. To make statistical significance comparable, we used the top 1256 genes in chimpanzees and three archaic humans as the inputs. The 1256 genes in chimpanzees and Altai Neanderthals are enriched in many common GO terms; however, “behaviour” (GO:0007610) and “cognition” (GO:0050890) have much higher significance in Altai Neanderthals than in chimpanzees (Figure 2). In addition, enrichment of such GO terms as “learning” (GO:0007612), “learning or memory” (GO:0007611), and “glucose homeostasis” (GO:0042593) is specific to Altai Neanderthals. Agreeing with the close relationship between Vindija Neanderthals and modern humans, the 134 genes in Vindija Neanderthals are enriched only in two GO terms with low significance (“smoothened signalling pathway” (GO:0007224) and “regulation of smoothened signalling pathway” (GO:0008589)) (Supplementary Table 7). These results help reveal the periods and positions of epigenetic regulation changes and help infer the periods of related phenotypic changes.

Earlier version: 1256 target genes whose DBSs have largest distances between modern humans and chimpanzees and Altai Neanderthals are enriched in different Biological Processes GO terms. Upon significance threshold = 0.05 (Benjamini-Hochberg FDR), the two sets of genes in chimpanzees and Altai Neanderthals are enriched in 199 and 409 GO terms (50 < terms size < 1000), respectively. Pink color indicates shared GO terms. Left: Top GO terms enriched in genes in chimpanzees. Right: Top GO terms enriched in genes in Altai Neanderthals.

Revised version: 1256 target genes whose DBSs have largest distances from modern humans to chimpanzees and Altai Neanderthals are enriched in different Biological Processes GO terms. Upon significance threshold = 0.05 (Benjamini-Hochberg FDR), the two gene sets in chimpanzees and Altai Neanderthals are enriched in 199 and 409 GO terms (50 < terms size < 1000), respectively. Pink color indicates shared GO terms. Left: Top GO terms enriched in genes in chimpanzees. Right: Top GO terms enriched in genes in Altai Neanderthals.

Earlier version: Gene regulation by HS lncRNAs shows differences in modern humans

Revised version: Epigenetic regulation by HS lncRNAs shows differences in modern humans

Earlier version: We next used the 1000 Human Genome data to analyze single nucleotide polymorphisms (SNPs with minimal allele frequency (MAF) ≥ 0.1) in DBSs (1000 Genomes Project Consortium et al., 2012). Many DBSs were found to be highly polymorphic (Supplementary Table 8). The most SNP-rich DBSs (> 10 SNPs, 5 times the average number of SNPs for all DBSs) are in genes important for human evolution, including IFNAR1, NFATC1, TAS1R3, INS, ST3GAL4, and FN3KRP (Table 2). These SNPs affect DBS binding affinity, no matter they are favored, neutral, or recessive. A feature of modern human diets is high sugar intake, which causes non-enzymatic oxidation of proteins (i.e., glycation) and deactivates protein functions. Among proteins encoded by genes with the most SNP-rich DBSs, TAS1R3 recognizes diverse natural and synthetic sweeteners, Insulin decreases blood glucose concentration, ST3GAL4 regulates protein glycosylation, and FN3KRP deglycates proteins. It worth noting that DBSs in these genes have counterparts with mostly changed sequences in chimpanzees and archaic humans (Table 2), indicating human-specific evolution of the regulation of these genes.

Revised version: To check whether the rewiring is still going on, we analysed single nucleotide polymorphisms (SNPs, with minimal allele frequency (MAF) ≥0.1) in DBSs using the 1000 Human Genome data (1000 Genomes Project Consortium et al., 2012). Many DBSs are highly polymorphic (>10 SNPs, 5 times the average number of SNPs for all DBSs); many DBSs have mostly changed sequence distances from modern humans to archaic humans and chimpanzees; and the intersection of the above two sets are in genes important for human evolution (Supplementary Table 8), including IFNAR1, NFATC1, TAS1R3, INS, ST3GAL4, and FN3KRP (Table 2). SNPs inside these DBSs, whether favored or recessive, affect DBS binding affinity in modern humans. A notable case is sugar metabolism. A feature of modern human diets is high sugar intake, which causes non-enzymatic oxidation of proteins (i.e., glycation) and deactivates protein functions. Among proteins encoded by genes whose DBSs are most SNP-rich and have mostly changed sequence distances, TAS1R3 recognizes diverse natural and synthetic sweeteners, Insulin decreases blood glucose concentration, ST3GAL4 regulates protein glycosylation, and FN3KRP deglycates proteins. DBSs in these genes not only have counterparts with mostly changed sequences in chimpanzees and archaic humans, but also are SNP-rich, indicating continuous evolution of these genes’ regulation by HS lncRNAs.

Earlier version: We further used multiple statistical tests, including XP-CLR (Chen et al., 2010), iSAFE (Akbari et al., 2018), Tajima’s D (Tajima, 1989), Fay-Wu’s H (Fay and Wu, 2000), the fixation index (Fst) (Weir and Cockerham, 1984), and linkage disequilibrium (LD) (Slatkin, 2008), to detect selection signals in HS lncRNAs and strong DBSs in the human population CEU, CHB, and YRI (the representative populations of Europeans, Asians, and Africans). These tests detected selection signals in several HS lncRNA genes in CEU and CHB but not in YRI and selection signals in more DBSs in CEU and CHB than in YRI (Supplementary Note 4, 5; Supplementary Table 9, 10, 11). Examples include the DBSs in the two pigmentation-related genes MC1R and MFSD12 and the two odor reception-related genes OR6C1 and TAS1R3, which may reflect the adaptation of Eurasians to new environments and lifestyles.

Revised version: We next used multiple statistical tests, including XP-CLR (Chen et al., 2010), iSAFE (Akbari et al., 2018), Tajima’s D (Tajima, 1989), Fay-Wu’s H (Fay and Wu, 2000), the fixation index (Fst) (Weir and Cockerham, 1984), and linkage disequilibrium (LD) (Slatkin, 2008), to detect selection signals in HS lncRNAs and strong DBSs in the modern human CEU, CHB, and YRI (the representative populations of Europeans, Asians, and Africans). These tests detected selection signals in several HS lncRNA genes in CEU and CHB but not in YRI and selection signals in more DBSs in CEU and CHB than in YRI (Supplementary Note 4, 5; Supplementary Table 9, 10, 11). Examples include the DBSs in the two pigmentation-related genes MC1R and MFSD12 and in the two odor reception-related genes OR6C1 and TAS1R3. The enriched SNPs, detected selection signals, and annotated gene functions suggest specific and ongoing influence of rewired gene expression on human evolution.

SNPs in DBSs have cis-effects on gene expression

Earlier version: To examine the impact of DBS sequence differences on gene expression, we assessed the expression quantitative trait loci (eQTLs) in DBSs using the Genotype-Tissue Expression (GTEx) data (GTEx Consortium, 2017). 1727 SNPs with MAF ≥ 0.1 in at least one population and absolute values of cis-effect size (ES) ≥ 0.5 in at least one tissue were identified in DBSs in autosomal genes (Supplementary Table 12,13). These eQTLs include 372 “conserved” ones (in DBSs in the three archaic and three modern humans) and 1020 “novel” ones (in DBSs only in the three modern humans). Many conserved eQTLs are in the brain and have high derived allele frequencies (DAFs) in all three modern populations. A notable eQTL that has high DAFs in all three populations and positive ES ≥ 0.5 in 44 out of 48 GTEx tissues is rs2246577 in the DBS in FN3KRP, whose encoded protein deglycates proteins to restore their function (Table 1). On the other hand, many novel eQTLs are in a particular tissue and have population-specific DAFs (Supplementary Note 6).

Revised version: To examine the impact of DBSs on gene expression, we assessed the expression quantitative trait loci (eQTLs) in DBS sequences using the Genotype-Tissue Expression (GTEx) data (GTEx Consortium, 2017). 1727 SNPs with MAF ≥0.1 in at least one population and with absolute values of cis-effect size (ES) ≥0.5 in at least one tissue were identified in DBSs in autosomal genes (Supplementary Table 12,13). These eQTLs include 372 “conserved” ones (also in DBSs in the three archaic humans) and 1020 “novel” ones (in DBSs only in modern humans). Many conserved eQTLs are in brain tissues and have high derived allele frequencies (DAFs) in all three modern populations. A notable eQTL that has high DAFs in all three populations and positive ES ≥0.5 in 44 out of 48 GTEx tissues is rs2246577 in the DBS in FN3KRP, whose encoded protein deglycates proteins to restore their function (Table 2). In contrast, many novel eQTLs are in a particular tissue and have population-specific DAFs (Supplementary Note 6).

Earlier version: We performed two analyses to examine that eQTLs in DBSs influence gene expression due to lncRNA-DBS interactions. First, we computed the expression correlation between HS lncRNAs and target transcripts with DBSs having eQTLs in specific tissues. Most (94%) HS lncRNA-target transcript pairs showed significant expression correlation in the tissues the eQTLs were identified (|Spearman’s rho| > 0.3 and FDR < 0.05). Second, we examined whether eQTLs are more enriched in DBSs than in promoter regions by computing and comparing the density of eQTLs in each DBS and each Ensembl-annotated promoter. The results confirm that eQTLs are more enriched in DBSs (one-sided Mann-Whitney test, p = 0.0) (Supplementary Note 6). Thus, many SNPs in DBSs have tissue-and population-specific cis-effects on gene expression.

Revised version: We performed two analyses to examine how eQTLs in DBSs influence gene expression via lncRNA-DBS interactions. First, we computed the expression correlation between HS lncRNAs and target transcripts with DBSs having eQTLs in specific tissues. Most (94%) HS lncRNA-target transcript pairs showed significant expression correlation in tissues in which the eQTLs were identified (|Spearman’s rho| >0.3 and FDR <0.05). Second, we examined whether eQTLs are more enriched in DBSs than in promoter regions by computing and comparing the density of eQTLs in each DBS and each Ensembl-annotated promoter. The results suggest that eQTLs are more enriched in DBSs than in promoters (one-sided Mann-Whitney test, p = 0.0) and that these eQTLs influence gene expression via lncRNA-DBS interactions (Supplementary Note 6). Tissue-specific eQTLs and population-specific SNPs in DBSs confirm the cis-effects of HS lncRNA-mediated epigenetic regulation on gene expression.

HS lncRNAs contributed to brain evolution from archaic to modern humans

Earlier version: We finally examined the impact of HS lncRNAs and their DBSs on gene expression in the GTEx tissues (GTEx Consortium, 2017). 40 autosomal HS lncRNAs are expressed robustly (median TPM > 0.1) in at least one tissue and form 198876 pairs with target transcripts in all tissues. 45% of these pairs show a significant expression correlation in specific tissues (Spearman’s |rho| > 0.3 and FDR < 0.05). In contrast, when randomly sampling 10000 pairs of lncRNAs and protein-coding transcripts genome-wide, the percent of pairs showing expression correlation with the same significance is only 2.3%. Notably, 56% of HS lncRNA-target transcript pairs with significant correlation are present in at least one brain region (Figure 3A), indicating substantial gene expression regulation by HS lncRNAs in the brain.

Revised version: We finally examined the impact of HS lncRNA-DBS interaction on gene expression in the GTEx tissues (GTEx Consortium, 2017). 40 autosomal HS lncRNAs are expressed robustly (median TPM > 0.1) in at least one tissue, and they form 198876 pairs with target transcripts in all tissues. 45% of these pairs show a significant expression correlation in specific tissues (Spearman’s |rho| >0.3 and FDR <0.05). In contrast, when randomly sampling 10000 pairs of lncRNAs and protein-coding transcripts genome-wide, the percent of pairs showing this level of expression correlation (Spearman’s |rho| >0.3 and FDR <0.05) is only 2.3%. Equally notable is that 56% of HS lncRNA-target transcript pairs with a significant correlation are present in at least one brain region (Figure 3A), indicating substantial gene expression regulation by HS lncRNAs in the brain.

Earlier version: The impact of HS lncRNAs on gene expression in GTEx tissues and organs. (A) Percentage distribution of correlated HS lncRNA-target transcript pairs across tissues and organs. Higher percentages of correlated pairs are in the brain than in other tissues and organs. (B) The distribution of significantly changed DBSs in HS lncRNA-target transcript pairs across tissues and organs between archaic and modern humans. Orange, red, and dark red indicate significant DBS changes from Denisovans (D), Altai Neanderthals and Denisovans (AD), and all of the three archaic humans (ADV). DBSs in HS lncRNA-target transcript pairs correlated in the seven brain regions have changed significantly and consistently since the Altai Neanderthals, Denisovans, and Vindija Neanderthals (one-sided two-sample Kolmogorov-Smirnov test, significant changes determined by FDR < 0.001).

Revised version: The impact of HS lncRNA-DBS interaction on gene expression in GTEx tissues and organs. (A) The percentage distribution of correlated HS lncRNA-target transcript pairs across GTEx tissues and organs. Higher percentages of correlated pairs are in the brain than in other tissues and organs. (B) The distribution of significantly changed DBSs in HS lncRNA-target transcript pairs across GTEx tissues and organs between archaic and modern humans. Orange, red, and dark red indicate significant DBS changes from Denisovans (D), Altai Neanderthals and Denisovans (A.D.), and all of the three archaic humans (ADV). DBSs in HS lncRNA-target transcript pairs correlated in seven brain regions (in dark red) have changed significantly and consistently since the Altai Neanderthals, Denisovans, and Vindija Neanderthals (one-sided two-sample Kolmogorov-Smirnov test, significant changes determined by FDR <0.001).

Earlier version: To verify the above conclusion, we analyzed nine experimental datasets (Supplementary Note 7). First, we analyzed two datasets of epigenetic studies (one examined H3K27ac and H3K4me2 profiling of human, macaque, and mouse corticogenesis, and the other examined gene expression and H3K27ac modification in eight brain regions in humans and four other primates) (Reilly et al., 2015; Xu et al., 2018). 84% and 73% of genes in the two datasets have DBSs for HS lncRNAs, significantly more enriched in HS lncRNA targets than the genome-wide background (p = 1.21e-21 and 1.2e-56, two-sided Fisher’s exact test). Second, we analyzed two datasets of PsychENCODE studies (one examined spatiotemporally differentially expressed genes and spatiotemporally differentially methylated sites in 16 brain regions, and the other examined the spatiotemporal transcriptomic divergence across human and macaque brain development) (Li et al., 2018; Zhu et al., 2018). In the first dataset, 65 HS lncRNAs were expressed in all 16 brain regions, 109 transcripts had spatiotemporally differentially methylated sites in their promoters, and 56 of the 109 transcripts had DBSs for HS lncRNAs. In the second dataset, 8951 genes showed differential expression between the human and macaque brain, and 72% of differentially expressed genes in the human brain have DBSs of HS lncRNAs. Thus, both datasets show significant enrichment for HS lncRNA regulation. Third, three studies identified several genes critically regulating cortical expansion (Florio et al., 2018; Johnson et al., 2018; Suzuki et al., 2018). Of the 40 protein-coding genes reported by these studies, 29 have DBSs of HS lncRNAs in their promoter regions (Supplementary Table 14). Thus, these genes are enriched with DBSs of HS lncRNAs compared with the genome-wide background (p < 0.01, two-sided Fisher’s exact test). Finally, we analyzed two datasets of brain organoids studies. By establishing and comparing cerebral organoids between humans, chimpanzees, and macaques, Pollen et al. identified 261 human-specific gene expression changes (Pollen et al., 2019). By fusing human and chimpanzee iPS cells and differentiating the hybrid iPS cells into hybrid cortical spheroids, Agoglia et al. generated a panel of tetraploid human-chimpanzee hybrid iPS cells and identified thousands of genes with divergent expression between humans and chimpanzees (Agoglia et al., 2021). The 261 and 1102 genes in the two datasets are enriched with DBSs of HS lncRNAs compared with the genome-wide background (p = 1.2e-16 and 3.4e-74).

Revised version: We analysed nine experimental datasets to verify the above conclusion substantial gene expression regulation by HS lncRNAs in the brain (Supplementary Note 7). First, we analysed two datasets of epigenetic studies (one examined H3K27ac and H3K4me2 profiling of human, macaque, and mouse corticogenesis, and the other examined gene expression and H3K27ac modification in eight brain regions in humans and four other primates) (Reilly et al., 2015; Xu et al., 2018). 84% and 73% of genes in the two datasets have DBSs for HS lncRNAs, significantly more enriched in HS lncRNA targets than the genome-wide background (p = 1.21e-21 and 1.2e-56, two-sided Fisher’s exact test). Second, we analysed two datasets of PsychENCODE studies (one examined spatiotemporally differentially expressed genes and spatiotemporally differentially methylated sites in 16 brain regions, and the other examined the spatiotemporal transcriptomic divergence across human and macaque brain development) (Li et al., 2018; Zhu et al., 2018). In the first dataset, 65 HS lncRNAs are expressed in all 16 brain regions, 109 transcripts have spatiotemporally differentially methylated sites in their promoters, and 56 of the 109 transcripts have DBSs for HS lncRNAs. In the second dataset, 8951 genes show differential expression between the human and macaque brains, and 72% of differentially expressed genes in the human brain have DBSs of HS lncRNAs. Thus, both datasets show significant enrichment for HS lncRNA regulation. Third, three studies identified several genes critically regulating cortical expansion (Florio et al., 2018; Johnson et al., 2018; Suzuki et al., 2018). Of the 40 protein-coding genes reported by these studies, 29 have DBSs of HS lncRNAs in their promoter regions (Supplementary Table 14). Thus, these genes are enriched with DBSs of HS lncRNAs compared with the genome-wide background (p < 0.01, two-sided Fisher’s exact test). Finally, we analysed two datasets of brain organoids studies. By establishing and comparing cerebral organoids between humans, chimpanzees, and macaques, Pollen et al. identified 261 human-specific gene expression changes (Pollen et al., 2019). By fusing human and chimpanzee iPS cells and differentiating the hybrid iPS cells into hybrid cortical spheroids, Agoglia et al. generated a panel of tetraploid human-chimpanzee hybrid iPS cells and identified thousands of genes with divergent expression between humans and chimpanzees (Agoglia et al., 2021). The 261 and 1102 genes in the two datasets are enriched with DBSs of HS lncRNAs compared with the genome-wide background (p = 1.2e-16 and 3.4e-74).

Earlier version: We next examined the evolution of the impact of HS lncRNAs and their DBSs on gene expression in the GTEx tissues. For each DBS in an HS lncRNA-target transcript pair correlated in a GTEx tissue, we computed the sequence distance of the DBS from modern humans to archaic humans. Then, we compared the distribution of DBS sequence distances in each tissue with the distribution of DBS sequence distances in all tissues (one-sided two-sample Kolmogorov-Smirnov test). DBSs in HS lncRNA-target transcript pairs correlated in brain regions have significantly changed sequence distances since the Altai Neanderthals and Denisovans (Figure 3B), suggesting that gene expression regulation by HS lncRNAs in the brain has experienced significant evolution and that HS lncRNAs contributed greatly to rapid brain evolution. Results from other GTEx tissues, including Kidney_Cortex, Ovary, and Skin, also suggest that the rewired gene expression has influenced human evolution.

Revised version: We next examined the evolution of the impact of HS lncRNAs and their DBSs on gene expression in the GTEx tissues. For each DBS in an HS lncRNA-target transcript pair correlated in a GTEx tissue, we computed the sequence distance of the DBS from modern humans to archaic humans. Then, we compared the distribution of DBS sequence distances in each tissue with the distribution of DBS sequence distances in all tissues (one-sided two-sample Kolmogorov-Smirnov test). DBSs in HS lncRNA-target transcript pairs correlated in brain regions have significantly changed sequence distances since the Altai Neanderthals and Denisovans (Figure 3B), suggesting that gene expression regulation by HS lncRNAs in the brain has experienced significant evolution and that HS lncRNAs contributed greatly to recent brain evolution (Ponce de Leon et al., 2021). Results from other GTEx tissues, including Kidney-Cortex, Ovary, and Skin-Sun Expressed (Figure 3B), also suggest continuous influences of rewired gene expression on human evolution.

Discussion

Earlier version: Modern humans differ from archaic humans and great apes little in genome sequences but greatly in phenotypes and behaviors, making what genomic differences critically determine the phenotypic and behavioral differences a compelling question. Several HS protein-coding genes promote the rapid expansion of neocortex (Fiddes et al., 2018; Florio et al., 2018; Pinson et al., 2022). Variations in TFBS cause gene expression changes (Franchini and Pollard, 2017). The rewiring of gene expression by TFs and TFBSs evolves (Johnson, 2017), so should the rewiring of gene expression by lncRNAs and their DBSs. However, the rewiring and evolution of the second case have been barely examined and poorly understood. This study identified HS lncRNAs and their DBSs, analyzed the rewiring and evolution of gene expression, and presents a new answer to the question of what genomic changes at what time critically drove human evolution. We performed seven CRISPR/Cas9 experiments to validate DBS prediction and analyzed multiple experimental datasets to validate our genome and transcriptome analyses. These results support the hypothesis that HS lncRNAs have distinctly rewired gene expression for adaptive evolution of humans.

Revised version: Modern humans differ from archaic humans and great apes trivially in genomes but greatly in phenotypes, making what genomic differences critically determine the phenotypic differences a profound question. This question has been addressed by many researchers along multiple approaches, and attention has been especially paid to new protein-coding genes, HARs, and evolved TF-TFBS interactions. Findings include HS protein-coding genes promote rapid neocortex expansion (Fiddes et al., 2018; Florio et al., 2018; Pinson et al., 2022), HARs are significantly enriched in domains important for human-specific 3D genome organization (Keough et al., 2023), and gene expression is rewired by TFs and TFBSs (Johnson, 2017). These findings suggest that species-specific lncRNAs and their DBSs may also substantially rewire gene expression. This study identified HS lncRNAs, their DBSs, and counterparts of these DBSs in ancient humans and multiple apes. Our analyses of the evolution of and interaction between HS lncRNAs and their DBSs generate evidence suggesting that HS lncRNAs and their DBSs have contributed greatly to human evolution by distinctly rewiring gene expression. Our analyses of multiple experimental datasets support this conclusion. The rewiring can be dated back to chimpanzees, has significantly evolved from archaic to modern humans, and still influences differences in modern humans. The occurrence of HS lncRNAs and their DBSs may have three situations – (a) HS lncRNAs preceded their DBSs, (b) HS lncRNAs and their DBSs co-occurred, (c) HS lncRNAs succeeded their DBSs. Our results support the third situation and the rewiring hypothesis.

Added in the revised version: We performed seven CRISPR/Cas9 experiments to validate DBS prediction and analysed multiple experimental datasets to support the main conclusion. Parameters and thresholds may influence data analysis results. When running Infernal, LongTarget, MAFFT7, and Phylip, default parameters were used. When analysing enrichment of GO terms in gene lists with strong and weak DBSs, the top and bottom 2000 genes were used, with the whole genome as the background. When analyzing enrichment of GO terms in gene lists with DBSs with large sequence distances, the threshold of 0.034 was used. DBSs in 4248, 1256, 2514, and 134 genes in chimpanzees, Altai Neanderthals, Denisovans, and Vindija Neanderthals have distance >0.034 to modern humans, and these numbers sufficiently characterize the human-chimpanzee differences and qualitatively agree with the phylogenetic relationships. Other parameter and threshold values would generate somewhat different results, but enriched GO terms with high significance will be quite stable. We note that most of the reported are findings rather than conclusions, indicative or suggestive of something that answer the primary question of what genomic differences critically determine the phenotypic differences between humans and apes and between modern and archaic humans.

Earlier version: Several findings are noteworthy. First, Neanderthal-inherited sequences have measurable impacts on gene expression, but the impacts are least detectable in the brain (McCoy et al., 2017). Here, we found that HS lncRNAs have more significantly influenced gene expression in the brain than in other tissues. Second, the regulation of genes by HS lncRNAs, making humans adapt to high sugar intake, has undergone great evolution. Third, genes regulated by HS lncRNAs are enriched in different functions at different stages. Fourth, old SNPs in DBSs are in genes regulating neural development, and new SNPs in DBSs are in genes regulating such processes as glucose metabolism.

Revised version: Several findings are noteworthy. First, Neanderthal-inherited sequences have measurable impacts on gene expression in modern humans, and the impacts are least detectable in the brain (McCoy et al., 2017). Here we find HS lncRNAs have more significantly influenced gene expression in the brain than in other tissues. Second, the regulation of many genes by HS lncRNAs has undergone great evolution, specific examples including those making humans adapt to high sugar intake. Third, intensively regulated genes are enriched in different functions at different periods of human evolution. Specifically, “old” SNPs in DBSs are in genes regulating neural development, and “new” SNPs in DBSs are in genes regulating such processes as glucose metabolism. These findings reveal the diverse, continuous, and subtle contribution of HS lncRNA to human evolution.

Earlier version: The above findings help reveal the functions of HS lncRNAs and the adaptive evolution of humans and highlight the important functions of other species-specific lncRNAs. They also raise three questions. First, the protein-coding regions in the mouse and human genomes are about 85% identical, but many mouse lncRNAs are rodent-specific (Yue et al., 2014). Thus, a question is how mouse-specific lncRNAs specifically rewire gene expression in mice and how the rewiring influences the transcriptional differences between humans and mice (Breschi et al., 2017; Hodge et al., 2019; Zhu et al., 2018). Second, whether mice and other mammals have a lncRNA like RP11-423H2.3 that substantially rewires gene expression species-specifically. Third, the extent clade- and species-specific lncRNAs would make speciation and evolution preordained. These questions are of great interest and have practical importance.

Revised version: The above findings not only reveal HS lncRNAs’ functions and adaptive human evolution, but also raise multiple new questions. First, the protein-coding regions in the mouse and human genomes are about 85% identical; however, many mouse lncRNAs are rodent-specific (Yue et al., 2014). Thus, two questions are how mouse-specific lncRNAs specifically rewire gene expression in mice and how human- and mouse-specific rewiring influences the cross-species transcriptional differences (Breschi et al., 2017; Hodge et al., 2019; Zhu et al., 2018). Second, whether mice and other mammals have lncRNAs such as RP11-423H2.3 that substantially rewires gene expression species-specifically. Third, to what extent clade- and species-specific lncRNAs would make speciation and evolution preordained. These interesting questions have practical importance.

Material and methods

Data sources

Earlier version: The sequences of human (GRCh38/hg38, GRCh37/hg19) and chimpanzee genomes (panTro5) were obtained from the UCSC Genome Browser (http://genome.UCSC.edu). Three high-quality archaic human genomes were obtained from the Max Planck Institute for Evolutionary Anthropology (https://www.eva.mpg.de/genetics/index.html), which include an Altai Neanderthal which lived approximately 122 thousand years ago (kya) (Pru fer et al., 2013), a Denisovan (an Asian relative of Neanderthals) which lived approximately 72 kya (Meyer et al., 2012), and a Vindija Neanderthal which lived approximately 52 kya (Prufer et al., 2017). The ancestral sequences for the human genome (GRCh37) (which were generated upon the EPO alignments of six primate species human, chimpanzee, gorilla, orangutan, macaque, and marmoset) were downloaded from the EBI website (ftp://ftp.1000genomes.ebi.ac.uk/vol1/ftp/phase1/analysis_results/supporting/ancestral_alignm ents/). The SNP data of modern humans were obtained from the 1000 Genomes Project (phase I) (1000 Genomes Project Consortium et al., 2012). The three modern human populations are CEU (Utah residents with Northern and Western European ancestry), CHB (Han Chinese in Beijing, China), and YRI (Yoruba in Ibadan, Nigeria), which contain 100, 99, and 97 individuals, respectively. The genetic map files of modern humans were downloaded from the 1000 Genomes Project website (ftp://ftp-trace.ncbi.nih.gov/1000genomes/ftp/technical/working/20110106_recombination_hotspots/). The 1000 Genomes phased genotype data in VCF format were downloaded from the International Genome Sample Resource website (http://www.internationalgenome.org/). Annotated human genes and transcripts were obtained from the Ensembl website (http://www.Ensembl.org). FASTA-format nucleotide sequences were converted from VCFs using the vcf-consensus program in the VCFtools package when necessary (Danecek et al., 2011). Data in the GeneCards Human Gene Database (www.genecards.org) were used to annotate genes and transcripts.

Revised version: The sequences of human (GRCh38/hg38, GRCh37/hg19) and chimpanzee genomes (panTro5) were obtained from the UCSC Genome Browser (http://genome.UCSC.edu). Three high-quality archaic human genomes were obtained from the Max Planck Institute for Evolutionary Anthropology (https://www.eva.mpg.de/genetics/index.html), which include an Altai Neanderthal which lived approximately 122 thousand years ago (kya) (Pru fer et al., 2013), a Denisovan (an Asian relative of Neanderthals) which lived approximately 72 kya (Meyer et al., 2012), and a Vindija Neanderthal which lived approximately 52 kya (Prufer et al., 2017). The ancestral sequences for the human genome (GRCh37) (which were generated upon the EPO alignments of six primate species human, chimpanzee, gorilla, orangutan, macaque, and marmoset) were downloaded from the EBI website (ftp://ftp.1000genomes.ebi.ac.uk/vol1/ftp/phase1/analysis_results/supporting/ancestral_alignments/). The SNP data of modern humans were obtained from the 1000 Genomes Project (phase I) (1000 Genomes Project Consortium et al., 2012). The three modern human populations are CEU (Utah residents with Northern and Western European ancestry), CHB (Han Chinese in Beijing, China), and YRI (Yoruba in Ibadan, Nigeria), which contain 100, 99, and 97 individuals, respectively. The genetic map files of modern humans were downloaded from the 1000 Genomes Project website (ftp://ftp-trace.ncbi.nih.gov/1000genomes/ftp/technical/working/20110106_recombination_hotspots/). The 1000 Genomes phased genotype data in VCF format were downloaded from the International Genome Sample Resource website (http://www.internationalgenome.org/). Annotated human genes and transcripts were obtained from the Ensembl website (http://www.Ensembl.org). FASTA-format nucleotide sequences were converted from VCFs using the vcf-consensus program in the VCFtools package when necessary (Danecek et al., 2011). Data in the GeneCards Human Gene Database (www.genecards.org) were used to annotate genes and transcripts.

Earlier version: RNA-sequencing (RNA-seq) data (transcripts per million [TPM]) and eQTL data of human tissues were obtained from the Genotype-Tissue Expression Project (GTEx, v8) website (https://gtexportal.org/) (GTEx Consortium, 2017). Histone modification and DNA methylation signals in multiple cell lines and ENCODE Candidate Cis-Regulatory Elements (cCRE) were obtained from the UCSC Genome Browser (http://genome.UCSC.edu). The genome-wide DNA binding sites of lncRNA NEAT1, MALAT1, and MEG3 were obtained from experimental studies (Mondal et al., 2015; West et al., 2014).

Revised version: RNA-sequencing (RNA-seq) data (transcripts per million [TPM]) and eQTL data of human tissues were obtained from the Genotype-Tissue Expression Project (GTEx, v8) website (https://gtexportal.org/) (GTEx Consortium, 2017). Histone modification and DNA methylation signals in multiple cell lines and ENCODE Candidate Cis-Regulatory Elements (cCRE) were obtained from the UCSC Genome Browser (http://genome.UCSC.edu). The genome-wide DNA binding sites of lncRNA NEAT1, MALAT1, and MEG3 were obtained from experimental studies (Mondal et al., 2015; West et al., 2014). The predicted and experimentally detected DBSs in target genes of MALAT1, NEAT1, and MEG3 are given in supplementary tables (Supplementary Table 15).

Identifying HS lncRNA genes

We used the Infernal program (Nawrocki et al., 2009), which searches orthologous RNA sequences upon sequence and structure alignment, to identify orthologous exons in 16 mammalian genomes for each exon in each of the 13562 GENCODE (v18)-annotated human lncRNA genes (Lin et al., 2019). The 16 mammals were chimpanzee (CSAC 2.1.4/panTro4), macaque (BGI CR_1.0/rheMac3), marmoset (WUGSC 3.2/calJac3), tarsier (Broad/tarSyr1), mouse lemur (Broad/micMur1), tree shrew (Broad/tupBel1), mouse (GRCm38/mm10), rat (Baylor3.4/rn4, RGSC6.0/rn6), guinea pig (Broad/cavPor3), rabbit (Broad/oryCun2), dog (Broad CanFam3.1/canFam3), cow (Baylor Btau_4.6.1/bosTau7), elephant(Broad/loxAfr3), hedgehog (EriEur2.0/eriEur2), opossum (Broad/monDom5), and platypus (WUGSC 5.0.1/ornAna1) (http://genome.UCSC.edu). If the number of orthologous exons of a human lncRNA gene in a genome exceeded half the exon number of the human lncRNA gene, these orthologous exons were assumed to form an orthologous lncRNA gene. If a human lncRNA gene had no orthologous gene in all of the 16 mammals, it was assumed to be human-specific.

Identifying DBSs of HS lncRNAs

Earlier version: LncRNAs bind to DNA sequences by forming RNA:DNA triplexes. Each triplex comprises triplex-forming oligonucleotides (TFO) in the lncRNA and a triplex-targeting site (TTS) in the DNA sequence. We used the LongTarget program to predict HS lncRNAs’ DNA binding domains (DBD) and binding sites (DBS) with the default parameters (Ruleset = all, TT penalty = -1000, CC penalty = 0, Offset = 15, Identity ≥ 60, Nt ≥ 50) (Lin et al., 2019; Wen et al., 2022). LongTarget program predicts DBDs and DBSs simultaneously, with a DBD comprising a set of densely overlapping TTOs and a DBS comprising a set of densely overlapping TTSs. For each HS lncRNAs, we predicted its DBSs in the 5000 bp promoter regions (-3500 bp upstream and +1500 bp downstream the transcription start site) of the 179128 Ensembl-annotated transcripts (release 79). For each DBS, a binding affinity was calculated by multiplexing the DBS length and the averaged Identity score of all TTSs (the Identity score is the percentage of paired nucleotides). Strong and weak DBSs were classified upon binding affinity ≥ 60 and < 60. A transcript whose promoter region contains a strong DBS of an HS lncRNA was assumed to be a target transcript of the HS lncRNA, and the gene containing this transcript was assumed to be a target gene of the HS lncRNA. As the 1000 Genomes Project (phase I) data and the archaic human genomes are based on GRCh37/hg19, the coordinates of DBSs were converted from GRCh38/hg38 to GRCh37/hg19 using the liftover program from the UCSC Genome Browser (Kuhn et al., 2013).

Revised version: LncRNAs bind to DNA sequences by forming RNA:DNA triplexes. Each triplex comprises triplex-forming oligonucleotides (TFO) in the lncRNA and a triplex-targeting site (TTS) in the DNA sequence. We used the LongTarget program to predict HS lncRNAs’ DNA binding domains (DBD) and binding sites (DBS) with the default parameters (Ruleset = all, T.T. penalty = −1000, CC penalty = 0, Offset = 15, Identity ≥ 60, Nt ≥ 50) (Lin et al., 2019; Wen et al., 2022). LongTarget program predicts DBDs and DBSs simultaneously, with a DBD comprising a set of densely overlapping TFOs and a DBS comprising a set of densely overlapping TTSs. For each HS lncRNA, we predicted its DBSs in the 5000 bp promoter regions (−3500 bp upstream and +1500 bp downstream the transcription start site) of the 179128 Ensembl-annotated transcripts (release 79). For each DBS, a binding affinity was calculated by multiplexing the DBS length and the averaged Identity score of all TTSs (the Identity score is the percentage of paired nucleotides). Strong and weak DBSs were classified upon binding affinity ≥ 60 and < 60. A transcript whose promoter region contains a strong DBS of an HS lncRNA was assumed to be a target transcript of the HS lncRNA, and the gene containing this transcript was assumed to be a target gene of the HS lncRNA. As the 1000 Genomes Project (phase I) data and the archaic human genomes are based on GRCh37/hg19, the coordinates of DBSs were converted from GRCh38/hg38 to GRCh37/hg19 using the liftover program from the UCSC Genome Browser (Kuhn et al., 2013).

Experimentally validating DBS prediction

Earlier version: A 157 bp sequence (chr17:80252565-80252721, hg19) containing the DBD of RP13-516M14.1, a 202 bp sequence (chr1:113392603-113392804, hg19) containing the DBD of RP11-426L16.8, and a 198 bp sequence (chr17:19460524-19460721, hg19) containing the DBD of SNORA59B, were knocked out in the HeLa cell line, RKO cell line, and SK-MES-1 cell line, respectively. Two sequences (chr1:156643524-156643684, chr10:52445649-52445740, hg38) containing the DBD of two wrongly transcribed noncoding sequences were knocked out in the HCT-116 and A549 cell lines, respectively. The seven cases of knocked out were performed by UBIGENE, Guangzhou, China, using CRISPR-UTM, a revised version of CRISPR/Cas9 technology (http://www.ubigene.com). The RNA-sequencing (RNA-seq) was performed by Novogene, Beijing, China (https://cn.novogene.com) and HaploX, Shenzhen, China (https://www.haplox.cn/), respectively.

Revised version: A 157 bp sequence (chr17:80252565-80252721, hg19) containing the DBD of RP13-516M14.1, a 202 bp sequence (chr1:113392603-113392804, hg19) containing the DBD of RP11-426L16.8, and a 198 bp sequence (chr17:19460524-19460721, hg19) containing the DBD of SNORA59B, were knocked out in the HeLa cell line, RKO cell line, and SK-MES-1 cell line, respectively. Two sequences (chr1:156643524-156643684, chr10:52445649-52445740, hg38) containing the DBD of two wrongly transcribed noncoding sequences were knocked out in the HCT-116 and A549 cell lines, respectively. The seven knockouts were performed by UBIGENE, Guangzhou, China (http://www.ubigene.com) using CRISPR-UTM, a revised version of CRISPR/Cas9 technology. Before and after the seven DBD knockouts, RNA-sequencing (RNA-seq) was performed by Novogene, Beijing, China (https://cn.novogene.com) and HaploX, Shenzhen, China (https://www.haplox.cn/). The reads were aligned to the human GRCh38 genome using the Hiasat2 program (Kim et al., 2019), and the resulting Sam files were converted to bam files using Samtools (Li et al., 2009). The Stringtie program was used to quantify gene expression levels (Pertea et al., 2015). Fold change of gene expression was computed using the edgeR package (Robinson et al., 2010), and significant up- and down-regulation of target genes after DBD knockout was determined upon |log2(fold change)| > 1 with FDR < 0.1.

Added in the revised version: Genome-wide DBSs of NEAT1, MALAT1, and MEG3 were experimentally detected (Mondal et al., 2015; West et al., 2014). We also used these data to validate DBS prediction by predicting DBSs of the three lncRNAs and checking the overlap between predicted and experimentally detected DBSs (Supplementary Table 15).

Mapping DBSs of HS lncRNAs in the chimpanzee and archaic human genomes

We used the liftover program from the UCSC Genome Browser (http://genome.UCSC.edu) to map loci of DBSs in the human genome hg38 to the chimpanzee genome Pan_tro 3.0 (panTro5). The mapping results were examined by checking the human-chimpanzee pairwise alignment (hg38 vs panTro5) in the UCSC Genome Browser (http://genome.UCSC.edu). 97.81% (100403/102651) of DBSs were mapped, and the 2248 unmapped DBSs in 429 genes were considered human-specific. We used vcf-consensus in the VCFtools package to extract the DBSs of HS lncRNAs from the VCF files of Altai Neanderthals, Denisovans, and Vindija Neanderthals. The variant with the highest quality score was selected whenever multiple variant calls were observed at a given locus. The obtained DBS sequences in the chimpanzees and three archaic humans are called counterparts of DBSs in these genomes.

Estimating sequence distances of DBSs between different genomes

We first aligned DBS sequences in the genomes of humans, chimpanzees, Altai Neanderthals, Denisovans, and Vindija Neanderthals using the MAFFT7 program for measuring sequence distances from modern humans to chimpanzees and archaic humans (Katoh and Standley, 2013). We then computed sequence distances using the dnadist program with the Kimura 2-parameter model in the PHYLIP (3.6) package (http://evolution.genetics.washington.edu/phylip.html) and using the Tamura-Nei model in the MEGA7 package (Kumar et al., 2016). The two methods generated equivalent results. The largest distance between DBSs in humans and their counterparts in chimpanzees is 5.383. Since 2248 DBSs in 429 human genes do not have counterparts in chimpanzees, we assumed that the sequence distances of these DBSs between humans and chimpanzees were 10.0. Since the top 20% of DBSs in chimpanzees (including these 2248 DBSs) have distances ≥ 0.034, we used 0.034 as the threshold to identify the “mostly changed” DBSs in the chimpanzees and three archaic humans.

We determined human ancestral sequences of DBSs upon the human ancestor sequences from the EBI website, which were generated upon the EPO alignments of six primate species. We used the above-mentioned methods to compute DBSs’ sequence distances from the human ancestor to chimpanzees, archaic humans, and modern humans. We found that when the human-chimpanzee ancestral sequence has the ancestral sequence (which means the inferring of ancestral allele is of high confidence), DBS distances from the human ancestor to modern humans are larger than to archaic humans, but this situation only accounts for about 63.8%. For many DBSs, the distances from the human ancestor to modern humans are smaller than to archaic humans (especially Neanderthals and Denisovans) and even to chimpanzees. This defect may be caused by the absence of archaic humans in building the human ancestral sequence.

Detecting positive selection signals in HS lncRNA genes and DBSs

We used multiple tests to detect positive selection signals in HS lncRNA genes. First, we used the XP-CLR test (parameters = -w1 0.001 300 100 -p0 0.95, window size = 0.1 cM, grid size = 100 bp) to perform six pairwise genome-wide scans (i.e., CEU-CHB, CEU-YRI, CHB-CEU, CHB-YRI, YRI-CEU, and YRI-CHB) (Chen et al., 2010). The upper 1% of scores across the entire genome in each pairwise scan were 34.6 in the CEU-YRI scan, 16.8 in the CEU-CHB scan, 45.0 in the CHB-YRI scan, 26.9 in the CHB-CEU scan, 14.1 in the YRI-CEU scan, and 14.1 in the YRI-CHB scan. These scores were used as the thresholds of positive selection signals in these populations. Second, we used the iSAFE program to scan each genomic region that contains a HS lncRNA gene and its 500-kb upstream and downstream sequences (Akbari et al., 2018). Strongly selected loci were detected only in CEU and CHB. Third, we used the VCFtools program to calculate Tajima’s D values for each HS lncRNA gene in CEU, CHB, and YRI (Danecek et al., 2011). The calculation was performed using a 1500-bp non-overlapping sliding window because the lengths of these genes exceed 1500 bp. To generate the background reference for judging the significant increase or decrease of Tajima’s D for HS lncRNA genes in a population, we calculated Tajima’s D values across the whole genome using a sliding window of 1500 bp. As values of Tajima’s D were compared with the background reference, a significant D<0 and D>0 indicate positive (or directional) selection and balancing selection, respectively, instead of population demography dynamics (Tajima, 1989). Fourth, we used the integrated Fst to detect positive selection signals in HS lncRNA genes. The Fst of each HS lncRNA gene was computed for three comparisons, i.e., CEU-YRI, CHB-YRI, and CHB-CEU. Extreme Fst values of SNPs were detected in HS lncRNA genes in the comparisons of CEU-YRI and CHB-YRI. Since allele frequencies for different loci vary in a genome and genetic drift may have different effects at different loci, we used a sliding window of 1500 for comparing Fst values of HS lncRNA genes with the genome-wide background. Extreme Fst values indicate positive selection. Finally, we applied linkage disequilibrium (LD) analysis to each HS lncRNA gene. We computed the pairwise LD (r²) in CEU, CHB, and YRI for common SNPs (with minimal allele frequency (MAF) ≥ 0.05 in at least one population) in HS lncRNA genes and DBSs using the PLINK program (Purcell et al., 2007). Significantly increased LD was detected in SNPs in HS lncRNA genes in CEU and CHB. The LD patterns were portrayed using the Haploview program (Barrett et al., 2005).

Earlier version: Next, we used the above tests to detect positive selection signals in DBSs. First, the 100 bp grid size of the XP-CLR test also allowed the detection of selection signals in DBSs. Second, we performed the Tajima’s D and Fay-Wu’s H tests (Fay and Wu, 2000; Tajima, 1989). We calculated Tajima’s D values for each DBS in CEU, CHB, and YRI using the VCFtools program (Danecek et al., 2011), with a sliding window of 147 bp (the mean length of strong DBSs). To generate the background reference for judging the significant increase or decrease of Tajima’s D in a population, we calculated Tajima’s D values across the whole genome using the same sliding window. As values of Tajima’s D were compared with the background reference, a significant D<0 and D>0 indicate positive (or directional) selection and balancing selection, respectively. Fay-Wu’s H values were calculated similarly using the VariScan program (Vilella et al., 2005). Calculating Fay-Wu’s H demands the ancestral sequences as the outgroup. We extracted the ancestral sequences of DBSs from the human ancestral sequence which was generated upon the EPO alignments of six primate species (human, chimpanzee, gorilla, orangutan, macaque, and marmoset) (ftp://ftp.1000genomes.ebi.ac.uk/vol1/ftp/phase1/analysis_results/supporting/ancestral_alignm ents/). Third, we computed the Fst to measure the frequency differences of alleles in DBSs between populations. For the CEU-YRI, CHB-YRI, and CHB-CEU pairwise comparisons, we used the revised VCFtools program to compute the weighted Fst for all SNPs in each DBS (Weir and Cockerham, 1984). Fourth, we integrated the weighted Fst values in the three populations into an “integrated Fst” which indicated whether the DBS locus was under selection in a certain population (Nielsen and Slatkin, 2013). We used sliding windows of 147 bp for comparing Fst values of DBSs with the genome-wide background. We empirically denoted the values of integrated Fst in the upper 10% of scores across the entire genome as statistically significant. To determine positive selection more reliably, we used Tajima’s D and the integrated Fst to jointly determine if a DBS was under positive selection in a population. The thresholds which determined the upper 10% of Tajima’s D values across the entire genome in CEU, CHB, and YRI were -0.97, -0.96, and -0.97, respectively, and the threshold which determined the upper 10% of integrated Fst values across the entire genome in the three populations was 0.22. For example, a DBS was assumed to be under positive selection in CEU if (a) the DBS had a Tajima’s D<-0.97 in CEU and Tajima’s D>0.0 in the two other populations and (b) the DBS had an integrated Fst>0.22. Analogously, a DBS was assumed to be under positive selection in both CEU and CHB if the DBS had a Tajima’s D <-0.97 in CEU, <-0.96 in CHB, and >0.0 in YRI, and had an integrated Fst>0.22.

Revised version: Next, we used the above tests to detect positive selection signals in DBSs. First, the 100 bp grid size of the XP-CLR test also allowed the detection of selection signals in DBSs. Second, we performed the Tajima’s D and Fay-Wu’s H tests (Fay and Wu, 2000; Tajima, 1989). We calculated Tajima’s D values for each DBS in CEU, CHB, and YRI using the VCFtools program (Danecek et al., 2011), with a sliding window of 147 bp (the mean length of strong DBSs). To generate the background reference for judging the significant increase or decrease of Tajima’s D in a population, we calculated Tajima’s D values across the whole genome using the same sliding window. As values of Tajima’s D were compared with the background reference, a significant D<0 and D>0 indicate positive (or directional) selection and balancing selection, respectively. Fay-Wu’s H values were calculated similarly using the VariScan program (Vilella et al., 2005). Calculating Fay-Wu’s H demands the ancestral sequences as the outgroup. We extracted the ancestral sequences of DBSs from the human ancestral sequence, which was generated upon the EPO alignments of six primate species (human, chimpanzee, gorilla, orangutan, macaque, and marmoset) (ftp://ftp.1000genomes.ebi.ac.uk/vol1/ftp/phase1/analysis_results/supporting/ancestral_alignm ents/). Third, we computed the Fst to measure the frequency differences of alleles in DBSs between populations. For the CEU-YRI, CHB-YRI, and CHB-CEU pairwise comparisons, we used the revised VCFtools program to compute the weighted Fst for all SNPs in each DBS (Weir and Cockerham, 1984). Fourth, we integrated the weighted Fst values in the three populations into an “integrated Fst” which indicated whether the DBS locus was under selection in a certain population (Nielsen and Slatkin, 2013). We used sliding windows of 147 bp for comparing Fst values of DBSs with the genome-wide background. We empirically denoted the values of integrated Fst in the upper 10% of scores across the entire genome as statistically significant. To determine positive selection more reliably, we used Tajima’s D and the integrated Fst to jointly determine if a DBS was under positive selection in a population. The thresholds that determined the upper 10% of Tajima’s D values across the entire genome in CEU, CHB, and YRI were −0.97, −0.96, and −0.97, respectively, and the threshold that determined the upper 10% of integrated Fst values across the entire genome in the three populations was 0.22. For example, a DBS was assumed to be under positive selection in CEU if (a) the DBS had a Tajima’s D<-0.97 in CEU and Tajima’s D>0.0 in the two other populations and (b) the DBS had an integrated Fst>0.22. Analogously, a DBS was assumed to be under positive selection in both CEU and CHB if the DBS had a Tajima’s D <-0.97 in CEU, <-0.96 in CHB, and >0.0 in YRI, and had an integrated Fst>0.22.

Functional enrichment analysis of genes

Earlier version: We used the g:Profiler program (with the parameter settings: Organism=Homo sapiens, Ordered query = No, Significance threshold = Benjamini-Hochberg FDR, User threshold = 0.05, 50 < terms size < 1000) and the Gene Ontology (GO) database to perform over-representation analysis (Raudvere et al., 2019). Over-representation analysis is used to determine which a priori defined gene sets are more present (over-represented) in a subset of “interesting” genes than what would be expected by chance. Here, we used the analysis to examine different target gene sets are enriched in what GO terms (Biological Process). The target gene sets included genes with strongest DBSs in humans, genes with weakest DBSs in humans, and genes with the mostly changed DBSs in chimpanzees and three archaic humans. To choose genes with strongest DBSs, we ranked genes in descending order based on DBS binding affinity and chose the top 2000 genes for over-representation analysis. To choose genes with weakest DBSs, we first ranked genes in ascending order based on DBS binding affinity, then ranked genes in descending order based on the number of regulatory HS lncRNAs, and finally chose the top 2000 genes with binding affinity ≤ 40 and regulatory HS lncRNAs ≥ 5. We used the distance threshold of 0.034 to choose genes with mostly changed DBSs. Genes with mostly changed DBSs in chimpanzees include the 429 genes in which 2248 DBSs lack counterparts. 4248, 1256, 2514, and 134 genes in chimpanzees, Altai Neanderthals, Denisovans, and Vindija Neanderthals have DBSs or their counterparts with distance ≥ 0.034. To make the statistical significance of enrichment analysis comparable across chimpanzees, Altai Neanderthals, and Denisovans, we used the same number of genes (the top 1256 genes whose DBSs have the largest DBS distances) as the input to the g:Profiler program.

Revised version: We used the g:Profiler program (with the parameter settings: Organism=Homo sapiens, Ordered query = No, Significance threshold = Benjamini-Hochberg FDR, User threshold = 0.05, 50 < terms size < 1000) and the Gene Ontology (GO) database to perform over-representation analysis (Raudvere et al., 2019). This analysis is used to determine what pre-defined gene sets (GO terms) are more present (over-represented) in a list of “interesting” genes than what would be expected by chance. The lists of genes included genes with strongest DBSs in humans, genes with weakest DBSs in humans, and genes with mostly changed DBSs in chimpanzees and three archaic humans. To choose genes with strong DBSs, we ranked genes in descending order based on binding affinity and chose the top 2000 genes. To choose genes with weak DBSs, we first ranked genes in ascending order based on binding affinity, then ranked genes in descending order based on the number of regulatory HS lncRNAs, and finally chose the top 2000 genes with binding affinity ≤40 and regulatory HS lncRNAs ≥5. We used the distance threshold of 0.034 to choose genes with mostly changed DBSs. Genes with mostly changed DBSs in chimpanzees include the 429 genes in which 2248 DBSs lack counterparts. 4248, 1256, 2514, and 134 genes in chimpanzees, Altai Neanderthals, Denisovans, and Vindija Neanderthals have DBSs with distance ≥0.034. To make the statistical significance comparable across chimpanzees, Altai Neanderthals, and Denisovans, we used the same number of genes (the top 1256 genes whose DBSs have the largest DBS distances) as the input to the g:Profiler program.

Analysing SNP frequencies in human populations

The frequency of common SNPs (MAF ≥0.05) in DBSs in the three modern human populations was computed using the VCFtools package (Danecek et al., 2011). The ancestral/derived states of SNPs were inferred upon the human ancestor sequences and were used to determine derived allele frequency.

Analysing the cis-effect of SNPs in DBSs on target gene expression

SNPs with MAF >0.1 in DBSs in any of the three modern human populations and absolute values of cis-effect size >0.5 (FDR <0.05) in any of the GTEx tissues were examined for an influence on the expression of the target genes (GTEx Consortium, 2017). SNPs that are eQTLs in the GTEx tissues and have biased derived allele frequencies (DAF) in the three modern human populations were examined to estimate whether the eQTL is population-specific.

Examining the tissue-specific impact of HS lncRNA-regulated gene expression

First, we examined the expression of HS lncRNA genes across the GTEx tissues. HS lncRNA genes with a median TPM value >0.1 in a tissue were considered robustly expressed in that tissue. Upon this criterion, 40 HS lncRNA genes were expressed in at least one tissue and were used to examine the impact of HS lncRNA regulation on gene and transcript expression (other HS lncRNAs may function in the cytoplasm). Since a HS lncRNA gene may have multiple transcripts, we adopted the transcript, which contains the predicted DNA binding domain (DBD) and has the largest TPM value as the representative transcript of the HS lncRNA. We calculated the pairwise Spearman’s correlation coefficient between the expression of a HS lncRNA (the representative transcript) and the expression of each of its target transcripts using the scipy.stats.spearmanr program in the scipy package. The expression of a HS lncRNA gene and a target transcript was considered to be significantly correlated if the |Spearman’s rho| > 0.3, with Benjamini-Hochberg FDR <0.05. We examined the percentage distribution of correlated HS lncRNA-target transcript pairs across GTEx tissues and organs (Figure 3A).

Earlier version: We examined all GTEx tissues to determine in which tissues HS lncRNA-regulated gene expression may have undergone significant changes from archaic to modern humans. If a pair of HS lncRNA and target transcript are robustly expressed in a tissue and their expression shows a significant correlation (|Spearman’s rho| > 0.3, with Benjamini-Hochberg FDR < 0.05) in the tissue, we computed the sequence distance of the HS lncRNA’s DBS in the transcript from the three archaic humans to modern humans. We compared the sequence distances of all DBSs in each tissue with the sequence distances of all DBSs in all GTEx tissues (i.e. the distribution of all DBS distances was taken as the background). A one-sided two-sample Kolmogorov-Smirnov test was used to examine whether the sequence distances of all DBSs in a specific tissue deviate from the background distribution (which reflects the “neutral evolution” of gene expression). For a tissue, if the test result was Benjamini-Hochberg FDR < 0.001, the tissue was considered to have significantly changed gene expression regulated by the HS lncRNA. We used different colors to mark tissues with significantly changed gene expression regulation since Altai Neanderthals, Denisovans, and Vindija Neanderthals, and used “D”, “AD”, and “ADV” to indicate changes since Denisovans, since Altai Neanderthals and Denisovans, and since Altai Neanderthals, Denisovans, and Vindija Neanderthals, respectively (Figure 3B).

Revised version: We examined all GTEx tissues to determine in which tissues HS lncRNA-regulated gene expression may have undergone significant changes from archaic to modern humans. If a pair of HS lncRNA and target transcript are robustly expressed in a tissue and their expression shows a significant correlation (|Spearman’s rho| >0.3, with Benjamini-Hochberg FDR <0.05) in the tissue, we computed the sequence distance of the HS lncRNA’s DBS in the transcript from the three archaic humans to modern humans. We compared the sequence distances of all DBSs in each tissue with the sequence distances of all DBSs in all GTEx tissues (i.e. the distribution of all DBS distances was taken as the background). A one-sided two-sample Kolmogorov-Smirnov test was used to examine whether the sequence distances of all DBSs in a specific tissue deviate from the background distribution (which reflects the “neutral evolution” of gene expression). For a tissue, if the test result was Benjamini-Hochberg FDR < 0.001, the tissue was considered to have significantly changed gene expression regulated by the HS lncRNA. We used different colors to mark tissues with significantly changed gene expression regulation since Altai Neanderthals, Denisovans, and Vindija Neanderthals, and used “D”, “A.D.”, and “ADV” to indicate changes since Denisovans, since Altai Neanderthals and Denisovans, and since Altai Neanderthals, Denisovans, and Vindija Neanderthals, respectively (Figure 3B).

Removed in the revised version: Declaration of competing interest

Removed in the revised version: The authors declare no competing interest.

Removed in the revised version: Author contributions

Removed in the revised version: H.Z. designed the study and drafted the manuscript. J.L. performed most data analyses. Y.W. and J.T. performed XP-CLR and iSAFE analyses. All authors have read the manuscript and consent to its publication.

Supporting information